Mitochondrial Protein Carboxyl-Terminal Alanine-Threonine Tailing Promotes Glioblastoma Tumor Growth by Regulating Mitochondrial Function

Author response:

The following is the authors’ response to the previous reviews.

eLife Assessment:

Glioblastoma is one of the most aggressive cancers without a cure. Glioblastoma cells are known to have high mitochondrial potential. This useful study demonstrates the critical role of the ribosome-associated quality control (RQC) pathway in regulating mitochondrial membrane potential and glioblastoma growth. Some assays are incomplete; further revision will improve the significance of this study.

For clarity, we propose revising the second sentence to: "It is well-established that certain cancer cells, such as glioblastoma cells, exhibit elevated mitochondrial membrane potential."

Reviewer #1 (Public Review):

Summary:

Cai et al have investigated the role of msiCAT-tailed mitochondrial proteins that frequently exist in glioblastoma stem cells. Overexpression of msiCAT-tailed mitochondrial ATP synthase F1 subunit alpha (ATP5) protein increases the mitochondrial membrane potential and blocks mitochondrial permeability transition pore formation/opening. These changes in mitochondrial properties provide resistance to staurosporine (STS)-induced apoptosis in GBM cells. Therefore, msiCAT-tailing can promote cell survival and migration, while genetic and pharmacological inhibition of msiCAT-tailing can prevent the overgrowth of GBM cells.

Strengths:

The CAT-tailing concept has not been explored in cancer settings. Therefore, the present provides new insights for widening the therapeutic avenue.

Your acknowledgment of our study's pioneering elements is greatly appreciated.

Weaknesses:

Although the paper does have strengths in principle, the weaknesses of the paper are that these strengths are not directly demonstrated. The conclusions of this paper are mostly well-supported by data, but some aspects of image acquisition and data analysis need to be clarified and extended.

We are grateful for your acknowledgment of our study’s innovative approach and its possible influence on cancer therapy. We sincerely appreciate your valuable feedback. In response, this updated manuscript presents substantial new findings that reinforce our central argument. Moreover, we have broadened our data analysis and interpretation, as well as refined our methodological descriptions.

Reviewer #2 (Public Review):

This work explores the connection between glioblastoma, mito-RQC, and msiCAT-tailing. They build upon previous work concluding that ATP5alpha is CAT-tailed and explore how CAT-tailing may affect cell physiology and sensitivity to chemotherapy. The authors conclude that when ATP5alpha is CAT-tailed, it either incorporates into the proton pump or aggregates and that these events dysregulate MPTP opening and mitochondrial membrane potential and that this regulates drug sensitivity. This work includes several intriguing and novel observations connecting cell physiology, RQC, and drug sensitivity. This is also the first time this reviewer has seen an investigation of how a CAT tail may specifically affect the function of a protein. However, some of the conclusions in this work are not well supported. This significantly weakens the work but can be addressed through further experiments or by weakening the text.

We appreciate the recognition of our study's novelty. To address your concerns about our conclusions, we have revised the manuscript. This revision includes new data and corrections of identified issues. Our detailed responses to your specific points are outlined below.

Reviewer #1 (Recommendations For The Authors):

(1) In Figure 1B, please replace the high-exposure blots of ATP5 and COX with representative results. The current results are difficult to interpret clearly. Additionally, it would be helpful if the author could explain the nature of the two different bands in NEMF and ANKZF1. Did the authors also examine other RQC factors and mitochondrial ETC proteins? I'm also curious to understand why CAT-tailing is specific to C-I30, ATP5, and COX-V, and why the authors did not show the significance of COX-V.

We appreciate your inquiry regarding the data. Additional attempts were made using new patient-derived samples; however, these results did not improve upon the existing ATP5⍺, (NDUS3)C-I30, and COX4 signals presented in the figure. This is possibly due to the fact that CAT-tail modified mitochondrial proteins represent only a small fraction of the total proteins in these cells. It is acknowledged that the small tails visible above the prominent main bands are not particularly distinct. To address this, the revised version includes updated images to better illustrate the differences. We believe the assertion that GBM/GSCs possess CAT-tailed proteins is substantiated by a combination of subsequent experimental findings. The figure (refer to new Fig. 1B) serves primarily as an introduction. It is important to note that the CAT-tailed ATP5⍺ plays a vital role in modulating mitochondrial potential and glioma phenotypes, a function which has been demonstrated through subsequent experiments.

It is acknowledged that the CAT-tail modification is not exclusive to the ATP5⍺protein. ATP5⍺ was selected as the primary focus of this study due to its prevalence in mitochondria and its specific involvement in cancer development, as noted by Chang YW et al. Future research will explore the possibility of CAT tails on other mitochondrial ETC proteins. Currently, NDUS3 (C-I30), ATP5⍺, and COX4 serve as examples confirming the existence of these modifications. It remains challenging to detect endogenous CAT-tailing, and bulk proteomics is not yet feasible for this purpose. COX4 is considered significant. We hypothesize that CAT-tailed COX4 may function similarly to the previously studied C-I30 (Wu Z, et al), potentially causing substantial mitochondrial proteostasis stress.

Concerning RQC proteins, our blotting analysis of GBM cell lines now includes additional RQC-related factors. The primary, more prominent bands (indicated by arrowheads) are, in our assessment, the intended bands for NEMF and ANKZF1. Subsequent blotting analyses showed only single bands for both ANKZF1 and NEMF, respectively. The additional, larger molecular weight band of NEMF, which was initially considered for property analysis (phosphorylation, ubiquitination, etc.), was not examined further as it did not appear in subsequent experiments (refer to new Fig. S1C).

References:

Chang YW, et al. Spatial and temporal dynamics of ATP synthase from mitochondria toward the cell surface. Communications biology. 2023;6(1).

Wu Z, et al. MISTERMINATE Mechanistically Links Mitochondrial Dysfunction With Proteostasis Failure. Molecular cell. 2019;75(4).

(2) In addition to Figure 1B, it would be interesting to explore CAT-tailed mETC proteins in cancer tissue samples.

This is an excellent point, and we appreciate the question. We conducted staining for ATP5⍺ and key RQC proteins in both tumor and normal mouse tissues. Notably, ATP5⍺ in GBM exhibited a greater tendency to form clustered punctate patterns compared to normal brain tissue, and not all of it co-localized with the mitochondrial marker TOM20 (refer to new Fig. S3C-E). Crucially, we observed a significant increase in NEMF expression within mouse xenograft tumor tissues, alongside a decrease in ANKZF1 expression (refer to new Fig. S1A, B). These findings align with our observations in human samples.

(3) Please knock down ATP5 in the patient's cells and check whether both the upper band and lower band of ATP5 have disappeared or not.

This control was essential and has been executed now. To validate the antibody's specificity, siRNA knockdown was performed. The simultaneous elimination of both upper and lower bands upon siRNA treatment (refer to new Fig. S2A) confirms they represent genuine signals recognized by the antibody.

(4) In Figure 1C and ID, add long exposure to spot aggregation and oligomer. Figure 1D, please add the blots where control and ATP5 are also shown in NHA and SF (similar to SVG and GSC827).

New data are included in the revised manuscript to address the queries. Specifically, the new Fig 1D now displays the full queue as requested, featuring blots for Control, ATP5α, AT3, and AT20. Our analysis reveals that AT20 aggregates exhibit higher expression and accumulation rates in GSC and SF cells.

Fig. 1C has been updated to include experimental groups treated with cycloheximide and sgNEMF. Our results show that sgNEMF effectively inhibits CAT-tailing in GBM cell lines, whereas cycloheximide has no impact. After consulting with the Reporter's original creator and optimizing expression conditions, we observed no significant aggregates with β-globin-non-stop protein, potentially due to the length of endogenous CAT-tail formation (as noted by Inada, 2020, in Cell Reports). Our analysis focused on the ratio of CAT-tailed (red box blots) and non-CAT-tailed proteins (green box blots). Comparing these ratios revealed that both anisomycin treatment and sgNEMF effectively hinder the CAT-tailing process, while cycloheximide has no effect.

(5) In Figure 1E, please double-check the results with the figure legend. ATP5A aggregated should be shown endogenously. The number of aggregates shown in the bar graph is not represented in micrographs. Please replace the images. For Figure 1E, to confirm the ATP5-specific aggregates, it would be better if the authors would show endogenous immunostaining of C-130 and Cox-IV.

Labels in Fig. 1E were corrected to reflect that the bar graph in Fig. 1F indicates the number of cells with aggregates, not the quantity of aggregates per cell. The presence

(6) Figure 3A. Please add representative images in the anisomycin sections. It is difficult to address the difference.

We appreciate your feedback. Upon re-examining the Calcein fluorescence intensity data in Fig. 3A, we believe the images accurately represent the statistical variations presented in Fig. 3B. To address your concerns more effectively, please specify which signals in Fig. 3A you find potentially misleading. We are prepared to revise or substitute those images accordingly.

(7) Figure 3D. If NEMF is overexpressed, is the CAT-tailing of ATP 5 reversed?

Thank you. Your prediction aligns with our findings. We've added data to the revised Fig. S6A, B, which demonstrates that both NEMF overexpression and ANKZF1 knockdown lead to elevated levels of CRC. This increase, however, was not statistically significant in GSC cells. A plausible explanation for this discrepancy is that the MPTP of GSC cells is already closed, thus any additional increase in CAT-tailing activity does not result in further amplification.

(8) Figure 3G. Why on the BN page are AT20 aggregates not the same as shown in Figure 2E?

We appreciate your inquiry regarding the ATP5⍺ blots, specifically those in the original Fig. 3G (left) and 2E (right). Careful observation of the ATP5⍺ band placement in these figures reveals a high degree of similarity. Notably, there are aggregates present at the top, and the diffuse signals extend downwards. Given that this is a gradient polyacrylamide native PAGE, the concentration diminishes towards the top. Consequently, the non-rigid nature of the Blue Native PAGE gel may lead to slight variations in the aggregate signals; however, the overall patterns are very much alike. To mitigate potential misinterpretations, we have rearranged the blot order in the new Fig. 3M.

(9) Figure 4D. The amount of aggregation mediated by AT20 is more compared to AT3. Why are there no such drastic effects observed between AT3 and AT20 in the Tunnel assay?

The previous Figure 4D presents the quantification of cell migration from the experiment depicted in Figure 4C. But this is a good point. TUNEL staining results are directly influenced by mitochondrial membrane potential and the state of mitochondrial permeability transition pores

(MPTP), not by the degree of protein aggregation. Our previous experiments showed comparable effects of AT3 and AT20 on mitochondria (Fig. 2E, 3K), which aligns with the expected similar outcomes on TUNEL staining. As for its biological nature, this could be very complicated. We hope to explore it in future studies.

(10) Figure 5C: The role of NEMF and ANKZF1 can be further clarified by conducting Annexin-PI assays using FACS. The inclusion of these additional data points will provide more robust evidence for CAT-tailing's role in cancer cells.

In response to your suggestion, we have incorporated additional data into the revised version.Using the Annexin-PI kit, we labeled apoptotic cells and detected them using flow cytometry (FACS). Our findings indicate that anisomycin pretreatment, NEMF knockdown (sgNEMF), and ANZKF1 upregulation (oeANKZF1) significantly increase the rate of STS-induced apoptosis compared to the control group (refer to new Fig. S9D-G).

(11) Figure 5F: STS is a known apoptosis inhibitor. Why it is not showing PARP cleavage? Also, cell death analysis would be more pronounced, if it could be shown at a later time point. What is the STS and Anisomycin at 24h or 48h time-point? Since PARP is cleaved, it would also be better if the authors could include caspase blots.

I guess what you meant to say here is "Staurosporine is a protein kinase inhibitor that can induce apoptosis in multiple mammalian cell lines." Our study observed PARP cleavage even in GSCs, which are typically more resistant to staurosporine-induced apoptosis (C-PARP in Fig. S9B). The ratio of C-PARP to total PARP increased. We selected a 180-minute treatment duration because longer treatments with STS + anisomycin led to a late stage of apoptosis and non-specific protein degradation (e.g., at 24 or 48 hours), making PARP comparisons less meaningful. Following your suggestion, we also examined caspase 3/7 activity in GSC cells treated with DMSO, CHX, and anisomycin. We found that anisomycin treatment also activated caspases (Fig. S9A).

(12) In Figure 5, the addition of an explanation, how CAT-tailing can induce cell death, would add more information such as BAX-BCL2 ratio, and cytochrome-c release from the mitochondria.

Thank you for your suggestion. In this study, we state that specific CAT-tails inhibit GSC cell death/apoptosis rather than inducing it. Therefore, we do not expect that examining BAX-BCL2 and mitochondrial cytochrome c release would offer additional insights.

(13) To confirm the STS resistance, it would be better if the author could do the experiments in the STS-resistant cell line and then perform the Anisomycin experiments.

Thank you. We should emphasize that our data primarily originates from GSC cells. These cells already exhibit STS-resistance when compared to the control cells (Fig. S8A-C).

(14) It would be more advantageous if the author could show ATP5 CATailed status under standard chemotherapy conditions in either cell lines or in vivo conditions.

This is an interesting question. It's worth exploring this question; however, GSC cells exhibit strong resistance to standard chemotherapy treatments like temozolomide (TMZ).

Additionally, we couldn't detect changes in CAT-tailed ATP5⍺ and thus did not include that data.

(15) In vivo (cancer mouse model or cancer fly model) data will add more weight to the story.

We appreciate your intriguing question. An effective approach would be to test the RQC pathway's function using the Drosophila Notch overexpression-induced brain tumor model. However, Khaket et al. have conducted similar studies, stating, "The RNAi of Clbn, VCP, and Listerin (Ltn), homologs of key components of the yeast RQC machinery, all attenuated NSC over-proliferation induced by Notch OE (Figs. 5A and S5A–D, G)." This data supports our theory, and we have incorporated it into the Discussion. While the mouse model more closely resembles the clinical setting, it is not covered by our current IACUC proposal. We intend to verify this hypothesis in a future study.

Reference:

Khaket TP, Rimal S, Wang X, Bhurtel S, Wu YC, Lu B. Ribosome stalling during c-myc translation presents actionable cancer cell vulnerability. PNAS Nexus. 2024 Aug 13;3(8):pgae321.

Reviewer #2 (Recommendations For The Authors):

Figure 1B, C: To demonstrate that Globin, ATP5alpha, and C-130 are CAT-tailed, it is necessary to show that the high mobility band disappears after NEMF deletion or mutagenesis of the NFACT domain of NEMF. This can be done in a cell line. The anisomycin experiment is not convincing because the intensity of the bands drops and because no control is done to show that the effects are not due to translation inhibition (e.g. cycloheximide, which inhibits translation but not CAT tailing). Establishing ATP5alpha as a bonafide RQC substrate and CAT-tailed protein is critical to the relevance of the rest of the paper.

Thank you for suggesting this crucial control experiment. To confirm the observed signal is indeed a bona fide CAT-tail, it's essential to demonstrate that NEMF is necessary for the CAT-tailing process. We have incorporated data from NEMF knockdown (sgNEMF) and cycloheximide treatment into the revised manuscript. Our findings show that both sgNEMF and anisomycin treatment effectively inhibit the formation of CAT-tailing signals on the reporter protein (Fig. 1C). Similarly, NEMF knockdown in a GSC cell line also effectively eliminated CAT-tails on overexpressed ATP5⍺ (Fig. S2B).

In general, the text should be weakened to reflect that conclusions were largely gleaned from artificial CAT tails made of AT repeats rather than endogenously CAT-tailed ATP5alpha. CAT tails could have other sequences or be made of pure alanine, as has been suggested by some studies.

Thank you for your reminder. We have reviewed the recent studies by Khan et al. and Chang et al., and we found their analysis of CAT tail components to be highly insightful. We concur with your suggestion regarding the design of the CAT tail sequence. We aimed to design a tail that maintained stability and resisted rapid degradation, regardless of its length. In the revised version, we clarify that our conclusions are based on artificial CAT tails, specifically those composed of AT repeat sequences (p. 9). We acknowledge that the presence of other sequence components may lead to different outcomes (p. 19).

Reference:

Khan D, Vinayak AA, Sitron CS, Brandman O. Mechanochemical forces regulate the composition and fate of stalled nascent chains. bioRxiv [Preprint]. 2024 Oct 14:2024.08.02.606406. Chang WD, Yoon MJ, Yeo KH, Choe YJ. Threonine-rich carboxyl-terminal extension drives aggregation of stalled polypeptides. Mol Cell. 2024 Nov 21;84(22):4334-4349.e7.

Throughout the work (e.g. 3B, C), anisomycin effects should be compared to those with cycloheximide to observe if the effects are specific to a CAT tail inhibitor rather than a translation inhibitor.

We agree that including cycloheximide control experiments is crucial. The revised version now incorporates new data, as depicted in Fig. S5A, B, illustrating alterations in the on/off state of MPTP following cycloheximide treatment. Furthermore, Fig. S6A, B present changes in Calcium Retention Capacity (CRC) under cycloheximide treatment. The consistency of results across these experiments, despite cycloheximide treatment, suggests that anisomycin's role is specifically as a CAT tail inhibitor, rather than a translation inhibitor.

Line 110, it is unclear what "short-tailed ATP5" is. Do you mean ATP5alpha-AT3? If so this needs to be introduced properly. Line 132: should say "may indicate accumulation of CAT-tailed protein" rather than "imply".

We acknowledge your points. We have clarified that the "short-tailed ATP5α" refers to ATP5α-AT3 and incorporated the requested changes into the revised manuscript.

Figure 1C: how big are those potential CAT-tails (need to be verified as mentioned earlier)?They look gigantic. Include a ladder.

In the revised Fig. 1D, molecular weight markers have been included to denote signal sizes. The aggregates in the previous Fig. 1C, also present in the control plasmid, are likely a result of signal overexposure. The CAT-tailed protein is observed just above the intended band in these blots. These aggregates have been re-presented in the updated figures, and their signal intensities quantified.

Line 170: "indicating that GBM cells have more capability to deal with protein aggregation". This logic is unclear. Please explain.

We appreciate your question and have thoroughly re-evaluated our conclusion. We offer several potential explanations for the data presented in Fig. 1D: (1) ATP5α-AT20 may demonstrate superior stability. (2) GSC (GBM) cells might lack adequate mechanisms to monitor protein accumulation. (3) GSC (GBM) cells could possess an increased adaptive capacity to the toxicity arising from protein accumulation. This discussion has been incorporated into the revised manuscript (lines 166-169).

Line 177: how do you know the endogenous ATP5alpha forms aggregates due to CAT-tailing? Need to measure in a NEMF hypomorph.

We understand your concern and have addressed it. Revised Fig. 3G, H demonstrates that a reduction in NEMF levels, achieved through sgNEMF in GSC cells, significantly diminishes ATP5α aggregation. This, in conjunction with the Anisomycin treatment data presented in revised Fig. 3E, F, confirms the substantial impact of the CAT-tailing process on this aggregation.

Line 218: really need a cycloheximide or NEMF hypomorph control to show this specific to CAT-tailing.

We have revised the manuscript to include data from sgNEMF and cycloheximide treatments, specifically Fig. 3G, H, and Fig. S5C, D, as detailed in our response above.

Lines 249,266, Figure 5A: The mentioned experiments would benefit from controls including an extension of ATP5alpha that was not alanine and threonine, perhaps a gly-ser linker, as well as an NEMF hypomorph.

We sincerely appreciate your insightful comments. In response, the revised manuscript now incorporates control data for ATP5α featuring a poly-glycine-serine (GS) tail. This data is specifically presented in Figs. S2E-G, S4E, S7A, D, E, and S8F, G. Our experimental findings consistently demonstrate that the overexpression of ATP5α, when modified with GS tails, had no discernible impact on protein aggregation, mitochondrial membrane potential, GSC cell mobility, or any other indicators assessed in our study.

Figure S5A should be part of the main figures and not in the supplement.

This has been moved to the main figure (Fig. 5C).

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: The rapid and sustained proliferation of cancer cells necessitates increased protein production, which, along with their disrupted metabolism, elevates the likelihood of translation errors. Ribosome-associated quality control (RQC), a recently identified mechanism, mitigates ribosome collisions resulting from frequent translation stalls. However, the precise pathophysiological role of the RQC pathway in oncogenesis remains ambiguous. Our research centered on the pathogenic implications of mitochondrial stress-induced protein carboxyl-terminal alanine and threonine tailing (msiCAT-tailing), a specific RQC response to translational arrest on the outer mitochondrial membrane, in glioblastoma (GBM). The presence of msiCAT-tailed mitochondrial proteins was observed commonly in glioblastoma stem cells (GSCs). Exogenous introduction of mitochondrial ATP synthase F1 subunit alpha (ATP5α) protein with artificial CAT-tail mimicking sequences enhanced mitochondrial membrane potential (ΔΨm) and inhibited mitochondrial permeability transition pore (MPTP) formation. These alterations in mitochondrial characteristics provided resistance to staurosporine (STS)-induced apoptosis in GBM cells. Consequently, msiCAT-tailing can foster cell survival and migration, whereas blocking msiCAT-tailing via genetic or pharmacological intervention can impede GBM cell overgrowth.

Revised version: The rapid and sustained proliferation of cancer cells necessitates increased protein production, which, along with their disrupted metabolism, elevates the likelihood of translation errors. Ribosome-associated quality control (RQC), a recently identified mechanism, mitigates ribosome collisions resulting from frequent translation stalls. However, the precise pathophysiological role of the RQC pathway in oncogenesis remains ambiguous. Our research centered on the pathogenic implications of mitochondrial stress-induced protein carboxyl-terminal alanine and threonine tailing (msiCAT-tailing), a specific RQC response to translational arrest on the outer mitochondrial membrane, in glioblastoma (GBM). The presence of msiCAT-tailed mitochondrial proteins was observed commonly in glioblastoma stem cells (GSCs). The exogenous introduction of the mitochondrial ATP synthase F1 subunit alpha (ATP5α) protein, accompanied by artificial CAT-tail mimicking sequences, enhanced mitochondrial membrane potential (ΔΨm) and inhibited the formation of the mitochondrial permeability transition pore (MPTP). These alterations in mitochondrial characteristics provided resistance to staurosporine (STS)-induced apoptosis in GBM cells. Consequently, msiCAT-tailing can foster cell survival and migration, whereas blocking msiCAT-tailing via genetic or pharmacological intervention can impede GBM cell overgrowth.

Introduction

Earlier version: Proteins are vital to biological processes, and their overproduction is crucial for rapidly proliferating cells, like cancer cells. To cope with this increased demand, cancer cells extensively reform the initiation, elongation, and termination phases of their protein synthesis (1). However, heightened protein translation elevates the chance of errors (2). Coupled with metabolic perturbations such as energy fluctuations and redox imbalances, the capacity to address disruptions during translation becomes indispensable. Ribosome-associated quality control (RQC) is a recently discovered suite of rescue mechanisms in eukaryotes that detect and resolve stalled, decelerated, or collided ribosomes during translation elongation or termination (3, 4).

Revised version: Proteins are vital to biological processes, and their overproduction is particularly crucial for rapidly proliferating cells, such as those found in cancer. To cope with this increased demand, cancer cells extensively reform the initiation, elongation, and termination phases of their protein synthesis (1). However, heightened protein translation elevates the chance of errors (2). Coupled with metabolic perturbations such as energy fluctuations and redox imbalances, the capacity to address disruptions during translation becomes indispensable. Ribosome-associated quality control (RQC) is a recently discovered suite of rescue mechanisms in eukaryotes that detect and resolve stalled, decelerated, or collided ribosomes during translation elongation or termination (3, 4).

Earlier version: RQC is a multi-step process initiated by the ZNF598/RACK1 complex recognizing the distinctive 40S-40S interface on collided ribosomes, which triggers ubiquitination of specific 40S subunit proteins (5, 6). Subsequently, the ASC-1 complex separates the leading ribosome (7, 8). Following this, events that transpire include: ribosomal subunits dissociation and recycling (9), modification of the nascent peptide chains by C-terminal alanine and threonine addition (CAT-tailing) (10), release of CAT-tailed products from the 60S subunits by ANKZF1/VMS1 (11), and degradation of aberrant peptides by the Ltn1/VCP/NEMF complex (4). The functional significance of CAT-tailed proteins produced during RQC remains incompletely understood. They may facilitate Ltn1-mediated ubiquitination (12) and promote the degradation of defective nascent peptides by exposing lysine residues (13, 14). Nonetheless, they are also prone to forming detergent-insoluble aggregates (15, 16). Furthermore, contingent upon the nature of the original protein and its subcellular location, CAT-tailed proteins might possess specific, albeit currently unclear, functions. Notably, CAT-tailed proteins have been implicated in the pathogenesis of several neurodegenerative diseases, indicating a significant role in their progression (17–19).

Revised version: RQC is a multi-step process initiated by the ZNF598/RACK1 complex, which recognizes the distinctive 40S-40S interface on collided ribosomes, triggering the ubiquitination of specific 40S subunit proteins (5, 6). Subsequently, the ASC-1 complex separates the leading ribosome (7, 8). Following this, events that transpire include: ribosomal subunit dissociation and recycling (9), modification of the nascent peptide chains by C-terminal alanine and threonine addition (CAT-tailing) (10), release of CAT-tailed products from the 60S subunits by ANKZF1/VMS1 (11), and degradation of aberrant peptides by the Ltn1/VCP/NEMF complex (4). The functional significance of CAT-tailed proteins produced during RQC remains incompletely understood. They may facilitate Ltn1-mediated ubiquitination (12) and promote the degradation of defective nascent peptides by exposing lysine residues (13, 14). Nonetheless, they are also prone to forming detergent-insoluble aggregates (15, 16). Furthermore, contingent upon the nature of the original protein and its subcellular location, CAT-tailed proteins might possess specific, albeit currently unclear, functions. Notably, CAT-tailed proteins have been implicated in the pathogenesis of several neurodegenerative diseases, indicating a significant role in their progression (17–19).

Earlier version: Cancerous cells demonstrate increased translation irregularities, including stop codon readthrough (20), frame-shifting (21), and oxidative stress-induced ribosomal arrest (22), suggesting a potential function for the RQC pathway. While CAT-tail modification on mitochondrial proteins due to compromised RQC has been noted in HeLa cells, the mechanistic involvement of RQC factors in cancer biology remains largely unexplored (17). Notably, the expression profile of various RQC factors (e.g., ASCC3, ABCE1, ANKZF1, and VCP) is dysregulated in cancer (23–26). Interestingly, RQC factors can display opposing functions in cancer development and suppression depending on specific circumstances, with some factors like ABCE1, ASCC3, and VCP suppressing cancer cell growth upon downregulation (23, 24, 26), while others like NEMF/Clbn and ZNF598 may promote it upon inhibition (27, 28). This indicates a nuanced, context-dependent role for RQC components in cancer cells, subject to both genetic and environmental modulators. A recent study examined ANKZF1’s mechanism in mitochondrial proteostasis and its influence on glioblastoma multiforme (GBM) progression (29). However, this study employed a non-physiological mitochondria-targeted GFP to induce matrix proteotoxicity, leaving the role of endogenous mitochondrial proteins in this process ambiguous.

Revised version: Cancerous cells exhibit increased translation irregularities, including stop codon readthrough (20), frame-shifting (21), and oxidative stress-induced ribosomal arrest (22), which suggests a potential role for the RQC pathway. While CAT-tail modification of mitochondrial proteins due to compromised RQC has been noted in HeLa cells, the mechanistic involvement of RQC factors in cancer biology remains largely unexplored (17). Notably, the expression profile of various RQC factors (e.g., ASCC3, ABCE1, ANKZF1, and VCP) is dysregulated in cancer (23–26). Interestingly, RQC factors can display opposing functions in cancer development and suppression depending on specific circumstances, with some factors like ABCE1, ASCC3, and VCP suppressing cancer cell growth upon downregulation (23, 24, 26), while others like NEMF/Clbn and ZNF598 may promote it upon inhibition (27, 28). This suggests a nuanced, context-dependent role for RQC components in cancer cells, influenced by both genetic and environmental factors. A recent study investigated the mechanism of ANKZF1 in mitochondrial proteostasis and its impact on glioblastoma multiforme (GBM) progression (29). However, this study employed a non-physiological mitochondria-targeted GFP to induce matrix proteotoxicity, leaving the role of endogenous mitochondrial proteins in this process ambiguous.

Earlier version: Mitochondrial stress leads to co-translational import anomalies, eliciting widespread CAT-tailing (mitochondrial stress-induced CAT-tail or msiCAT-tail) of nuclear-encoded mitochondrial proteins, including C-I30 (Complex-I 30 kDa subunit protein, NDUS3) (17, 30). The functional ramifications of these msiCAT-tailed proteins in mitochondrial biology remain poorly elucidated. Given that CAT-tailing imparts new properties to proteins, it may contribute to distinctive features of cancer cell mitochondria, such as hyperpolarization (31, 32) and resistance to drug-induced apoptosis linked to high mitochondrial membrane potential (ΔΨm) (33–35). This membrane potential across the inner membrane of mitochondria, essential for ATP production by OXPHOS, is sustained by the electron transport chain (Complexes I to IV), which pumps protons (H+) into the intermembrane space (36), and ATP synthase (Complex V), which leverages this gradient (37). While numerous malignant cells exhibit reduced OXPHOS despite high energy demands (38), the mechanisms by which they maintain or elevate ΔΨm remain an unresolved question (31).

Revised version: Mitochondrial stress leads to co-translational import anomalies, eliciting widespread CAT-tailing (mitochondrial stress-induced CAT-tail or msiCAT-tail) of nuclear-encoded mitochondrial proteins, including C-I30 (Complex-I 30 kDa subunit protein, NDUS3) (17, 30). The functional ramifications of these msiCAT-tailed proteins in mitochondrial biology remain poorly elucidated. Given that CAT-tailing imparts new properties to proteins, it may contribute to the distinctive features of cancer cell mitochondria, such as hyperpolarization (31, 32) and resistance to drug-induced apoptosis linked to a high mitochondrial membrane potential (Δψm) (33–35). This membrane potential across the inner membrane of mitochondria, essential for ATP production by OXPHOS, is sustained by the electron transport chain (Complexes I to IV), which pumps protons (H+) into the intermembrane space (36), and ATP synthase (Complex V), which leverages this gradient (37). While numerous malignant cells exhibit reduced OXPHOS despite high energy demands (38), the mechanisms by which they maintain or elevate ΔΨm remain an unresolved question (31).

In this study, we investigated msiCAT-tailing modification on the mitochondrial ATP synthase F1 subunit alpha (ATP5α). We discerned that msiCAT-tailed ATP5α is present in GBM. The mimic short-tailed ATP5α (ATP5α-AT3 in subsequent studies) can integrate into the ATP synthase, leading to an augmented ΔΨm and attenuated mitochondrial permeability transition pore (MPTP) assembly and opening. Consequently, msiCAT-tailed ATP5α enhances GBM cell resistance to programmed cell death induced by staurosporine (STS) and temozolomide (TMZ), thereby fostering cancer cell survival, proliferation, and migration. Conversely, impeding msiCAT-tailing diminishes cancer cell growth and resensitizes GBM cells to apoptosis. Our findings underscore the involvement of CAT-tailed mitochondrial proteins in tumorigenesis and emphasize the significance of the RQC pathway in oncobiology. These outcomes suggest that components and products of the RQC pathway may offer promising therapeutic targets for GBM.

Results

Presence of msiCAT-tailed Proteins in Glioblastoma Cells

Earlier version: While dysregulation of individual ribosome-associated quality control (RQC) factors is documented across various cancers (e.g., adenocarcinoma, non-small cell lung, prostate, and colon carcinomas), a comprehensive analysis of the RQC pathway in glioblastoma (GBM) has been lacking (23–26). Our analysis of transcriptomic data from a cohort of 153 GBM patients and 206 healthy controls, sourced from public datasets, revealed significantly elevated expression (logFC (fold change) > 1; adj.P.Val < 0.001) of RQC pathway genes, such as ABCE1, ASCC1-3, RACK1, and VCP, in GBM cells(39). Conversely, ANKZF1 was significantly downregulated (logFC = −0.43, adj.P.Val = 0.0005) (Fig. 1A, Table S1). The expression change in these genes implies RQC pathway activation and potential accumulation of CAT-tailed proteins in GBM. Mitochondrial stress-induced protein mitochondrial Complex-I 30 kDa (C-I 30, also known as NDUS3), an endogenous RQC substrate with msiCAT-tails, was previously identified in HeLa cells (17). Examination of patient-derived GBM stem cells (GSCs) and normal neural stem cells (NSCs) revealed that GSCs, unlike NSCs, exhibited several msiCAT-tailed mitochondrial proteins, including NDUS3, COX4 (Cytochrome c Oxidase subunit 4), and ATP5α (ATP synthase F1 subunit alpha). Consistent with the detection of these msiCAT-tailing signals, increased NEMF (Nuclear Export Mediator Factor) levels (10) and decreased ANKZF1 (Ankyrin Repeat and Zinc-finger Peptidyl tRNA Hydrolase 1) expression (11) were observed in patient-derived GSCs (Fig. 1B), further indicative of enhanced CAT-tailing activation, mirroring bioinformatics findings in GBM samples. A murine GBM model exhibited analogous RQC pathway alterations, with increased NEMF and decreased ANKZF1 expression in transplanted SB28 gliomas compared to normal brain tissue (Fig. S1A, B).

Revised version: While dysregulation of individual ribosome-associated quality control (RQC) factors is documented across various cancers (e.g., adenocarcinoma, non-small cell lung, prostate, and colon carcinomas), a comprehensive analysis of the RQC pathway in glioblastoma (GBM) has been lacking (23–26). Our analysis of transcriptomic data from a cohort of 153 GBM patients and 206 healthy controls, sourced from public datasets, revealed significantly elevated expression (logFC (fold change) > 1; adj.P.Val < 0.001) of RQC pathway genes, such as ABCE1, ASCC1-3, RACK1, and VCP, in GBM cells(39). Conversely, ANKZF1 was significantly downregulated (logFC = -0.43, adj.P.Val = 0.0005) (Fig. 1A, Table 1). The expression change in these genes implies activation of the RQC pathway and potential accumulation of CAT-tailed proteins in GBM. Mitochondrial stress-induced protein mitochondrial Complex-I 30 kDa (C-I 30, also known as NDUS3), an endogenous RQC substrate with msiCAT-tails, was previously identified in HeLa cells (17). Examination of patient-derived GBM stem cells (GSCs) and normal neural stem cells (NSCs) revealed that GSCs, unlike NSCs, exhibited several msiCAT-tailed mitochondrial proteins, including NDUS3, COX4 (Cytochrome c Oxidase subunit 4), and ATP5α (ATP synthase F1 subunit alpha). Consistent with the detection of these msiCAT-tailing signals, increased NEMF (Nuclear Export Mediator Factor) levels (10) and decreased ANKZF1 (Ankyrin Repeat and Zinc-finger Peptidyl tRNA Hydrolase 1) expression (11) were observed in patient-derived GSCs (Fig. 1B), further indicative of enhanced CAT-tailing activation, mirroring bioinformatics findings in GBM samples. A murine GBM model exhibited analogous RQC pathway alterations, with increased NEMF and decreased ANKZF1 expression in transplanted SB28 gliomas compared to normal brain tissue (Fig. S1A, B).

Evidence for msiCAT-tailing on mitochondrial proteins in GBM cells

(A) RQC gene expression levels in GBM tumor tissues (n=153) compared to normal brain tissues (n=206) (unpaired Student’s t-test; *, logFC (fold change)>1; adj.P.Val<0.001). (B) Western blot analysis of msiCAT-tailed mitochondrial proteins and RQC factors in patient-derived GSC and control NSC cells, using ACTIN as the loading control. Red arrowheads indicate short CAT-tailed mitochondrial proteins; “short” and “long” refer to exposure time; the red numbers represent fold changes compared to controls (NSC). (C) Western blot of 5×FLAG-tagged β-globin reporter proteins in GBM and control cells, showing more CAT-tailed proteins in GBM cells, using ACTIN as the loading control. The red numbers represent fold changes compared to controls (NHA without any treatment); the purple numbers represent the ratio of red (CAT-tailed) to green (non-CAT-tailed) sections. (D) Western blot of overexpressed ATP5α-AT3 and ATP5α-AT20 in GBM and control cells, using GAPDH as the loading control; arrowheads indicate endogenous ATP5α, ATP5α-AT3, ATP5α-AT20, and oligomers/aggregates of msiCAT-tailed ATP5α proteins. The purple numbers represent the ratio of red (exogenous) to green (endogenous) sections. (E) Immunofluorescence staining shows endogenous ATP5α protein aggregates in GBM cells, with TOM20 (red) as a mitochondrial marker. White arrows indicate ATP5α protein aggregates. (F) Quantification of E (n=3; chi-squared test; ***, P < 0.001; ****, P < 0.0001); the total number of cells counted is indicated in the columns.

The subsequent experiments were conducted using two GBM cell lines, SF268 (SF in figures) (40) and GSC827 (GSC in figures) (41), and two control cell lines, SVG p12 (SVG in figures) and Normal Human Astrocytes E6/E7/hTERT (NHA in figures) (42). RQC protein expression analysis revealed decreased ANKZF1 and increased ABCE1, ASCC3, and NEMF expression in GSC827 and SF268 cells, consistent with findings in patient-derived GSCs (Fig. S1C). Intriguingly, induction of CAT-tailing on a Flag-tagged β-globin reporter via a non-stop protein translation system demonstrated significantly higher CAT-tailed protein (β-globin-nonstop) production in GBM cells (43). This process was inhibitable by the CAT-tailing elongation inhibitor anisomycin and NEMF knockdown (sgNEMF), but not cycloheximide treatment, as evidenced by a decreased ratio of CAT-tailed (Red) to non-CAT-tailed bands (Green) (Fig. 1C).

Next, to investigate the biological implications of CAT-tailing, artificial CAT-tails were introduced to mitochondrial proteins. Due to the variability in CAT-tailing, prior research simulated this process by adding alanine-threonine (AT) repeat tails to the C-terminus of mitochondrial proteins (17). According to recent studies, the chosen tail sequence can be stabilized by its high threonine content (44, 45). ATP5α, a highly abundant mitochondrial protein with roles in cancer, was selected to study the unique functions of CAT-tailed forms (46, 47). siATP5α knockdown first confirmed the upper band signal in GSCs as authentic ATP5α, demonstrated by its disappearance concurrent with the main band’s weakening (Fig. S2A). Then, we confirmed that this upper band signal corresponded to changes in CAT-tailing, which could be effectively inhibited by NEMF knockdown and anisomycin treatment (Fig. S2B). Due to the indistinct nature of the endogenous msiCAT-tailed ATP5α signal, exogenously expressed Flag-ATP5α was utilized here.

To investigate the potential new function provided by CAT-tailed proteins, control (SVG and NHA) and GBM (SF268 and GSC827) cell lines overexpressed ATP5α with three (ATP5α-AT3) or twenty (ATP5α-AT20) AT repeats. Consistent with earlier findings, only the long-tailed ATP5α-AT20 exhibited post-translational modifications and detergent-resistant insoluble aggregates, appearing as slower migrating bands and a high molecular weight smear in protein electrophoresis (Fig. 1D). Based on comparing exogenously expressed (indicated by red boxes) to endogenous proteins (indicated by green boxes), GBM cell lines (GSC827, SF268) showed increased accumulation of ATP5α-AT20 compared to control cells (SVG, NHA). This accumulation may occur due to increased stability and reduced degradation of long-tailed proteins, a malfunctioning protein quality control system, enhanced cellular tolerance to protein accumulation, or a combination of these factors. Subcellular localization analysis showed that the short AT tail (AT3) did not significantly alter ATP5α’s mitochondrial localization, similar to the tailless protein. However, a significant portion of ATP5α-AT20 was found in the cytoplasm near mitochondria, forming protein aggregates, with the highest proportion in highly malignant GSC cells (Fig. S2C, D). Notably, poly-glycine-serine tails (short, GS3, and long, GS20) did not induce insoluble protein aggregation or intracellular punctate distribution (Fig. S2E-G), highlighting the importance of specific amino acid composition.

Importantly, in GBM cells, both exogenous tailed proteins and the endogenous ATP5α formed clusters attached to the outer mitochondrial membrane (Fig. 1E, F). Similar aggregate formation in GBM cells was also observed with other mitochondrial proteins, such as NDUS3 (Fig. S3A, B). Furthermore, we examined the mouse GBM models. Akin to in vitro culture, ATP5α in transplanted SB28 glioma formed more punctate signals and did not always colocalize with the mitochondrial marker TOM20 (Fig. S3C-E). These findings collectively indicate a disruption of the RQC pathway, leading to the presence of msiCAT-tailed proteins in GBM cells.

msiCAT-tailed ATP5α Elevates Mitochondrial Membrane Potential (ΔΨm)

Some cancer cells exhibit altered mitochondrial physiology, maintaining or increasing mitochondrial membrane potential (ΔΨm) despite reduced respiration. This was observed in patient-derived GSC cells, which demonstrated higher ΔΨm but lower ATP production than control NSC cells (Fig. 2A, B). Similarly, GBM cell lines, GSC827 and SF268, displayed comparable or higher ΔΨm and lower ATP levels relative to the control NHA cell line (42) (Fig. S4A-C). Genetic inhibition of msiCAT-tailing, via NEMF knockdown (sgNEMF) or ANZKF1 overexpression (oeANZKF1) (Fig. S4D), as well as pharmacological inhibition by anisomycin treatment, effectively reduced ΔΨm in GBM cells but not in NHA cells (Fig. 2C, D).

Impact of msiCAT-tailed ATP5IZ proteins on mitochondrial functions in GBM cells

(A) TMRM staining shows a high mitochondrial membrane potential in patient-derived GSC cells (n=3; unpaired Student’s t-test; ***, P < 0.001; ****, P < 0.0001). (B) ATP measurement shows a low mitochondrial ATP production in patient-derived GSC cells (n=3; unpaired Student’s t-test; **, P < 0.01; ***, P < 0.001). (C, D) JC-10 staining reveals a reduced mitochondrial membrane potential in GBM cells, but not in NHA control cells, upon both genetic (C) and pharmacological (D) inhibition of the msiCAT-tailing pathway (n=3; unpaired Student’s t-test; ***, P < 0.001; ****, P < 0.0001; ns, not significant). (E) JC-10 staining reveals an increased mitochondrial membrane potential in GBM cells, but not in control cells, upon overexpression of ATP511-AT3 and ATP511-AT20 (n=3; unpaired Student’s t-test; ****, P < 0.0001; ns, not significant). (F) Western blot of FLAG-tagged ATP511, NEMF, and ANKZF1 in GBM cells and control cells, using ACTIN as the loading control. (G) JC-10 staining reveals an increased mitochondrial membrane potential in GBM cells, but not in NHA control cells, upon overexpression of ATP511-AT3 and ATP511-AT20 with concurrent genetic inhibition of the endogenous msiCAT-tailing pathway (n=3; unpaired Student’s t-test; *, P < 0.05; **, P < 0.01). (H) BN-PAGE western blot of ATP511 and Flag shows that ATP511-AT3 is incorporated into the mitochondrial Complex-V (ATP synthase), while ATP511-AT20 forms high molecular weight protein aggregates in GBM cells. SC: respiratory supercomplex; C-V: Complex-V/ATP synthase. (I, K) Oxygen consumption rate (OCR) data indicate a reduction in mitochondrial oxygen consumption in SF268 cells expressing ATP511-AT3 and ATP511-AT20. Oligomycin (1.5 µM), FCCP (1.0 µM), and rotenone/antimycin A (R/A, 0.5 µM) were sequentially added. (J, L) Statistics of mitochondrial respiration parameters in (I, K), including non-mitochondrial respiration, basal respiration, maximum respiration, spare respiration, proton leaks, and ATP production (n=3; unpaired Student’s t-test; *, P < 0.05; **, P < 0.01; ***, P < 0.001; ****, P < 0.0001; ns, not significant).

Earlier version: Our next investigation of msiCAT tail proteins revealed their impact on mitochondrial function. Expression of Flag-tagged ATP511-AT3 and ATP511-AT20 in GBM and control cell lines elevated ΔΨm specifically in GBM cells (Fig. 2E). Overexpression of ATP511-GS3 and ATP511-GS20 did not exhibit this effect (Fig. S4E). To our surprise, even with suppressed endogenous CAT-tailing through sgNEMF and oeANZKF1 in GSC cells, the introduced AT3 and AT20 proteins could still effectively elevate ΔΨm (Fig. 2F, G). This finding suggests that CAT- tailing of ATP511 may be a significant contributor to the observed mitochondrial phenotype (Fig. 2G). Blue Native Polyacrylamide Gel Electrophoresis (BN-PAGE) illustrated distinct effects based on CAT-tail length. ATP511-AT3 integrated into the mitochondrial respiratory chain complex, whereas ATP511-AT20 formed high molecular weight complexes or remained as monomers (Fig. 2H). In mitochondrial physiological activity assays using the Agilent Cell Mitochondrial Stress Test, the oxygen consumption rate (OCR) was directly measured to assess mitochondrial respiration. Our findings indicate that expressing both ATP511-AT3 and ATP511-AT20 negatively impacted mitochondrial oxidative phosphorylation. This impairment leads to a reduction in ATP synthesis, basal respiration, and maximal respiration rates (Fig. 2I-L). These data suggest that both short and long tails on ATP511 proteins influence mitochondrial function, although potentially through different mechanisms. Short CAT-tails may directly act on ATP synthase function and thus affect the respiratory chain complex, while long CAT-tails form protein aggregates, causing mitochondrial proteostasis stress and thus indirectly affecting mitochondrial respiration (17, 29). This differential impact of CAT-tail length suggests a nuanced regulation of mitochondrial function mediated by ATP511 modifications.

Revised version: Our next investigation of msiCAT tail proteins revealed their impact on mitochondrial function. Expression of Flag-tagged ATP51111-AT3 and ATP511-AT20 in GBM and control cell lines elevated ΔΨm specifically in GBM cells (Fig. 2E). Overexpression of ATP511-GS3 and ATP511-GS20 did not exhibit this effect (Fig. S4E). To our surprise, even with suppressed endogenous CAT-tailing through sgNEMF and oeANZKF1 in GSC cells, the introduced AT3 and AT20 proteins could still effectively elevate ΔΨm (Fig. 2F, G). This finding suggests that CAT-tailing of ATP511 may be a significant contributor to the observed mitochondrial phenotype (Fig. 2G). Blue Native Polyacrylamide Gel Electrophoresis (BN-PAGE) illustrated distinct effects based on CAT-tail length. ATP511-AT3 integrated into the mitochondrial respiratory chain complex, whereas ATP511-AT20 formed high molecular weight complexes or remained as monomers (Fig. 2H). In mitochondrial physiological activity assays using the Agilent Cell Mitochondrial Stress Test, the oxygen consumption rate (OCR) was directly measured to assess mitochondrial respiration. Our findings indicate that expressing both ATP511-AT3 and ATP511-AT20 negatively impacted mitochondrial oxidative phosphorylation. This impairment leads to a reduction in ATP synthesis, basal respiration, and maximal respiration rates (Fig. 2I-L). These data suggest that both short and long tails on ATP511 proteins influence mitochondrial function, although potentially through different mechanisms. Short CAT-tails may directly act on ATP synthase function and thus affect the respiratory chain complex, while long CAT-tails form protein aggregates, causing mitochondrial proteostasis stress and thus indirectly affecting mitochondrial respiration (17, 29). This differential impact of CAT-tail length suggests a nuanced regulation of mitochondrial function mediated by ATP511 modifications.

msiCAT-tailing Influences Mitochondrial Permeability Transition Pore (MPTP) Dynamics

Beyond its traditionally recognized role in ATP production, the F1F0 ATP synthase has garnered increasing attention as a potential structural component of the mitochondrial permeability transition pore (MPTP) complex (48–50). Given the possibility that CAT-tailed proteins like ATP511 might modulate MPTP function, this investigation sought to elucidate the mechanism by which msiCAT-tailing modulates MPTP dynamics (open-close state). Comparative analyses conducted in GBM and control cells revealed that MPTP in GSC827 cells predominantly exists in a closed conformation, indicated by strong Calcein signals. Notably, the treatment of anisomycin, a pharmacological CAT-tailing inhibitor, effectively induced MPTP opening in GSC827 cells, as indicated by decreased Calcein signals (Fig. 3A, B). This effect was concomitant with the diminished aggregation of endogenous ATP5α (Fig. 3E, F). Furthermore, corroborative evidence was obtained through genetic manipulation. Specifically, genetic inhibition of CAT-tailing via NEMF knockdown (sgNEMF) resulted in a similar decrease in Calcein signaling and a reduction in ATP5α accumulation (Fig. 3C, D, G, H), aligning with the results obtained using anisomycin. In contrast, treatment with cycloheximide, a general translation inhibitor, did not significantly alter Calcein or ATP511 aggregation signals (Fig. S5A-D), suggesting that non-specific translation inhibition does not impact the mitochondrial MPTP state. The crucial role of CAT-tail modifications on ATP5α in modulating MPTP status was further substantiated by the observation that overexpression of artificially synthesized AT repeat tails (AT3 and AT20) restored Calcein signals despite the inhibition of endogenous CAT-tailing (Fig. S5E).

msiCAT-tailing product regulates MPTP status in GBM cells

Earlier version: (A, C) MPTP activity assay shows reduced mitochondrial permeability transition pore opening in GSC cells compared to NHA (control) cells. Pharmacological (A, anisomycin 200 nM) and genetic (sgNEMF) inhibition of CAT-tailing reverse it. (B, D) Quantification of (A, C) (n=3; unpaired Student’s t-test; ****, P < 0.0001; ns, not significant). (E, G) Immunofluorescence staining shows that anisomycin treatment (E) and sgNEMF (G) inhibit endogenous ATP511 protein aggregation in GBM cells, using TOM20 (red) as a mitochondrial marker. (F, H) Quantification of (E, G) (n=3; chi-squared test; *, P < 0.05; ***, P < 0.001); the total number of cells counted is indicated in the columns. (I) Calcium retention capacity (CRC) assay of isolated mitochondria, measured with Calcium Green-5N dye, shows high CRC in GBM cells compared to control NHA cells. CsA (Cyclosporin A, MPTP inhibitor) serves as a positive control. (J) Statistic of (I) shows attenuated CRC in mitochondria pretreated with anisomycin or with sgNEMF (n=3; unpaired Student’s t-test; ***, P < 0.001; ****, P < 0.0001). (K) MPTP activity assay shows that ectopic expression of ATP511-AT3 and ATP511-AT20 inhibits MPTP opening in GBM cells. (L) Quantification of (K) (n=3; unpaired Student’s t-test; ****, P < 0.0001). (M) BN-PAGE western blot shows that ATP511-AT3 and ATP511-AT20 expression alters ANT1/2 protein patterns in GBM cells, resulting in a missing band (circled in yellow dashed line), and formation of high molecular weight aggregates. SC: respiratory supercomplex; C-V: Complex V/ATP synthase.\

Revised version: (A, C) MPTP activity assay shows reduced mitochondrial permeability transition pore opening in GSC cells compared to NHA (control) cells. Pharmacological (A, anisomycin 200 nM) and genetic (sgNEMF) inhibition of CAT-tailing reverse it. (B, D) Quantification of (A, C) (n=3; unpaired Student’s t-test; ****, P < 0.0001; ns, not significant). (E, G) Immunofluorescence staining shows that anisomycin treatment (E) and sgNEMF (G) inhibit endogenous ATP511 protein aggregation in GBM cells, using TOM20 (red) as a mitochondrial marker. (F, H) Quantification of (E, G) (n=3; chi-squared test; *, P < 0.05; **, P < 0.01; ***, P < 0.001); the total number of cells counted is indicated in the columns. (I) The calcium retention capacity (CRC) assay of isolated mitochondria, measured using the Calcium Green-5N dye, reveals a significantly higher CRC in GBM cells compared to control NHA cells. CsA (Cyclosporin A, MPTP inhibitor) serves as a positive control. (J) Statistic of (I) shows attenuated CRC in mitochondria pretreated with anisomycin or with sgNEMF (n=10; unpaired Student’s t-test; ***, P < 0.001; ****, P < 0.0001). (K) MPTP activity assay shows that ectopic expression of ATP511-AT3 and ATP511-AT20 inhibits MPTP opening in GBM cells. (L) Quantification of (K) (n=3; unpaired Student’s t-test; ****, P < 0.0001). (M) BN-PAGE western blot shows that ATP511-AT3 and ATP511-AT20 expression alters ANT1/2 protein patterns in GBM cells, resulting in a missing band (circled in yellow dashed line), and formation of high molecular weight aggregates. SC: respiratory supercomplex; C-V: Complex V/ATP synthase.

Earlier version: The MPTP is recognized to participate in the transient efflux of protons, calcium ions (Ca²), and other signaling molecules from the mitochondrial matrix during brief opening episodes (51). To quantitatively evaluate the MPTP open/closed state, the mitochondrial Ca² Retention Capacity (CRC) assay was employed, which measures the amount of Ca² required to elicit MPTP opening. Our results revealed that GSC827 cells exhibited a greater CRC than NHA cells. Pre-treatment with anisomycin or knockdown of NEMF (sgNEMF) significantly decreased the CRC in GBM cells, indicating MPTP opening upon the loss of CAT-tailed proteins (Fig. 3I, J). Consistent with Calcein staining results (Fig. S5A, B), cycloheximide treatment did not substantially alter CRC measurements (Fig. S6A, B). Conversely, enhancing CAT-tailing (e.g., via oeNEMF and siANKZF1) led to an increase in CRC (Fig. S6A, B), although this effect was less pronounced in GSCs, potentially due to their inherently active CAT-tailing and closed MPTP.

Revised version: The MPTP is recognized to participate in the transient efflux of protons, calcium ions (Ca2+), and other signaling molecules from the mitochondrial matrix during brief opening episodes (51). To quantitatively evaluate the MPTP open/closed state, the mitochondrial Ca2 Retention Capacity (CRC) assay was employed, which measures the amount of Ca2+ required to elicit MPTP opening. Our results revealed that GSC827 cells exhibited a greater CRC value than NHA cells. Pre-treatment with anisomycin or knockdown of NEMF (sgNEMF) significantly decreased the CRC in GBM cells, indicating MPTP opening upon the loss of CAT-tailed proteins (Fig. 3I, J). Consistent with Calcein staining results (Fig. S5A, B), cycloheximide treatment did not substantially alter CRC measurements (Fig. S6A, B). Conversely, enhancing CAT-tailing (e.g., via oeNEMF and siANKZF1) led to an increase in CRC (Fig. S6A, B), although this effect was less pronounced in GSCs, potentially due to their inherently active CAT-tailing and closed MPTP.

To further investigate the impact of specific AT repeat tails on MPTP opening, artificial AT repeat tails on ATP511 were introduced into GBM cells. It was found that the short AT tail (AT3) inhibited MPTP opening, while the long AT tail (AT20) displayed a weaker effect (Fig. 3K, L), potentially due to their different integration into ATP synthase (Fig. 2G). Complex co-immunoprecipitation assay did not detect direct interactions between ATP511 with AT3 or AT20 tails and MPTP components Cyclophilin D (Cyp-D) and adenine nucleotide translocator 2 (ANT2) (Fig. S6C). However, Cyp-D expression was reduced upon ectopic expression of ATP511-AT3 and ATP511-AT20, suggesting decreased MPTP formation (Fig. S6D). Intriguingly, BN-PAGE analysis revealed that both ATP511-AT3 and ATP511-AT20 altered ANT1/2-containing complexes, with expected bands disappearing (indicated by *) and aggregates forming (at the top), supporting the notion that ATP synthase is integrated into the MPTP supercomplex due to the spatial proximity of the ANT1/2 complex and ATP synthase (Fig. 3M). In conclusion, msiCAT-tailed ATP511 proteins, particularly those with short AT3 tails, are integrated into ATP synthase and have a substantial influence on modulating MPTP status.

msiCAT-Tailing Boosts GBM Cell Migration and Resistance to Apoptosis

The elevated mitochondrial membrane potential (ΔΨm) and constricted MPTP resulting from msiCAT-tailed ATP511 and other mitochondrial proteins may enhance cellular stress resilience. We first investigated how the msiCAT-tailing mechanism affects GBM cells at the cellular level. MTT assays (52) revealed that overexpressing short (AT3) and long (AT20) AT repeat tails, fused to ATP511, significantly improved GBM cell viability, but not that of NHA cells (Fig. 4A, B). However, short (GS3) and long (GS20) GS repeat tails did not affect GBM cell viability (Fig. S7A). In addition, in vitro transwell migration assays (53) and wound healing assays (54) showed that GBM cells overexpressing AT repeat-tailed ATP5α exhibited increased cell invasion and accelerated wound healing, indicating enhanced cell migration (Fig. 4C, D, S7B, C). Notably, neither ATP5α alone nor GS repeat-tailed proteins showed comparable changes (Fig. S7D, E). Furthermore, overexpressing AT3- and AT20-tailed proteins effectively conferred phenotypes associated with increased GBM cell activity, such as enhanced survival and migration, even with inhibited endogenous CAT-tailing machinery activity (e.g., sgNEMF and oeANKZF1) (Fig. 4E-G). It is worth noting that ANKZF1 knockdown in U87 and U251 cell lines can cause aberrant mitoGFP accumulation, possibly reducing cellular adaptability (29), suggesting varying mitochondrial adaptability to proteostasis stress across cell lines. Supporting this, initial experiments showed that mild expression of ATP5α-AT3 and ATP5α-AT20 did not induce strong mitochondrial proteotoxic responses, as evidenced by the lack of significant upregulation in LONP1, mtHSP70, and HSP60 mRNA levels (Fig. S7F).

Earlier version: msiCAT-tailed ATP5IZ protein promotes GBM progression

Revised version: msiCAT-tailed ATP511 protein promotes GBM progression

Earlier version: (A) MTT assay indicates increased proliferation caused by ATP511-AT3 and ATP511-AT20 expression in GBM cells (n=3; unpaired Student’s t-test; **, P < 0.01; ***, P < 0.001). (B) MTT assay indicates no change in proliferation caused by ATP511-AT3 and ATP511-AT20 expression in NHA cells (n=3; unpaired Student’s t-test; ns, not significant). (C) Transwell assay reveals enhanced migration induced by ATP511-AT3 and ATP511-AT20 expression in GBM (SF) cells but not in control (NHA) cells. (D) Quantification of (C) shows the number of migrated cells (n=3; unpaired Student’s t-test; ****, P < 0.0001; ns, not significant). (E) MTT assay indicates an increased proliferation in GBM cells, upon overexpression of ATP511-AT3 and ATP5-AT20 with concurrent genetic inhibition of the endogenous msiCAT-tailing pathway (n=3; unpaired Student’s t-test; *, P < 0.05; **, P < 0.01). (F) Transwell assay reveals enhanced migration upon overexpression of ATP511-AT3 and ATP511-AT20 with concurrent genetic inhibition of the endogenous msiCAT-tailing pathway. (G) Quantification of (F) shows the number of migrated cells (n=3; unpaired Student’s t-test; ***, P < 0.001; ****, P < 0.0001). (H) TUNEL staining shows that staurosporine (STS, 1 µM)-induced apoptosis is attenuated by ATP511-AT3 and ATP511-AT20 expression in GBM cells, using TUNEL-Cy3 as an apoptotic cell indicator and DAPI as a nucleus indicator. (I) Quantification of (H) shows the percentage of TUNEL-positive cells in the population (n=3; unpaired student’s t-test; ***, P < 0.001), using DMSO as the vehicle control. (J) MTT assay indicates an enhanced resistance to temozolomide (TMZ, 150 µM) induced by ATP511-AT3 and ATP511-AT20 expression. The TMZ-treated/SF-Ctrl group is used as the control (n=3; unpaired Student’s t-test; ***, P < 0.001).

Revised version: (A) MTT assay indicates increased proliferation caused by ATP511-AT3 and ATP511-AT20 expression in GBM cells (n=3; unpaired Student’s t-test; **, P < 0.01; ***, P < 0.001). (B) MTT assay indicates no change in proliferation caused by ATP511-AT3 and ATP511-AT20 expression in NHA cells (n=3; unpaired Student’s t-test; ns, not significant). (C) Transwell assay reveals enhanced migration induced by ATP511-AT3 and ATP511-AT20 expression in GBM (SF) cells but not in control (NHA) cells. (D) Quantification of (C) shows the number of migrated cells (n=3; unpaired Student’s t-test; ***, P < 0.001; ns, not significant). (E) MTT assay indicates an increased proliferation in GBM cells, upon overexpression of ATP511-AT3 and ATP5-AT20 with concurrent genetic inhibition of the endogenous msiCAT-tailing pathway (n=3; unpaired Student’s t-test; *, P < 0.05; **, P < 0.01). (F) Transwell assay reveals enhanced migration upon overexpression of ATP511-AT3 and ATP511-AT20 with concurrent genetic inhibition of the endogenous msiCAT-tailing pathway. (G) Quantification of (F) shows the number of migrated cells (n=3; unpaired Student’s t-test; ***, P < 0.001; ****, P < 0.0001). (H) TUNEL staining shows that staurosporine (STS, 1 µM)-induced apoptosis is attenuated by ATP511-AT3 and ATP511-AT20 expression in GBM cells, using TUNEL-Cy3 as an apoptotic cell indicator and DAPI as a nucleus indicator. (I) Quantification of (H) shows the percentage of TUNEL-positive cells in the population (n=3; unpaired student’s t-test; ***, P < 0.001), using DMSO as the vehicle control. (J) MTT assay indicates an enhanced resistance to temozolomide (TMZ, 150 µM) induced by ATP511-AT3 and ATP511-AT20 expression. The TMZ-treated/SF-Ctrl group is used as the control (n=3; unpaired Student’s t-test; ***, P < 0.001).

GBM cells exhibit increased resistance to staurosporine (STS)-induced apoptosis, supported by fewer TUNEL-positive cells (Fig. S8A, B) and markedly diminished PARP-1 (Poly ADP-ribose polymerase) cleavage (Fig. S8C), a marker of AIF-mediated apoptosis (55). To investigate the role of CAT-tailed ATP5α proteins in this resistance, we overexpressed proteins with mimetic tails in GBM cells. Overexpression of both short tail (ATP5α-AT3) and long tail (ATP5α-AT20) significantly enhanced resistance to STS-induced apoptosis, as shown by TUNEL staining (Fig. 4H, I) and flow cytometry (Fig. S8D, E), indicating a strong link between protein CAT-tailing and tumorigenesis. In contrast, control short (GS3) and long (GS20) GS tails failed to confer such resistance (Fig. S8F, G). Consistent with these findings, overexpression of artificial CAT-tailed ATP5α proteins also increased the resistance of GBM cells to temozolomide (TMZ)-induced apoptosis (Fig. 4J). Taken together, these results suggest that RQC-induced CAT-tailing on ATP5α protein plays a role in GBM resistance to drug-induced apoptosis.

Earlier version: GBM Cell Progression is Hindered by RQC Pathway Inhibition

Revised version: RQC Pathway Inhibition hinders GBM Cell Progression

Prior research indicates the RQC pathway’s mediated msiCAT-tailing plays an important role in GBM progression, suggesting it as a potential therapeutic target. To explore this, patient-derived Glioblastoma Stem Cell (GSC) lines were treated with anisomycin, an inhibitor of CAT-tailing. GSC lines displayed higher sensitivity to anisomycin than normal neural stem cells (NSCs) (Fig. 5A). Similarly, genetic inhibition of the RQC pathway via NEMF knockdown (sgNEMF) or ANKZF1 overexpression (oeANZKF1) in the SF268 GBM cell line also suppressed GBM growth (Fig. 5B). Notably, control NHA cell proliferation was also inhibited by these genetic changes, indicating the broad significance of NEMF and ANKZF1 in cell proliferation (Fig. 5C). The RQC pathway appears to have a more pronounced effect on GBM cell migration. In in vitro transwell assays, sgNEMF or oeANZKF1 notably decreased GBM cell migration without affecting NHA cells (Fig. 5D, E). Consistently, anisomycin treatment impaired GSC cell migration, but not NHA cell migration (Fig. 5F, G).

Inhibition of msiCAT-tailing impedes GBM progression

Earlier version: (A) Cell viability assay shows greater sensitivity to anisomycin treatment in patient-derived GSC cells than control NSC cells at 48 hours (n=3; unpaired Student’s t-test; ***, P < 0.001; ****, P <0.0001; compared to controls at the corresponding dose). (B) MTT assay indicates reduced GBM cell proliferation by genetic inhibition of the msiCAT-tailing pathway (n=3; unpaired Student’s t-test; **, P < 0.01; ***, P < 0.001; ****, P <0.0001, compared to controls at the corresponding time). (C) MTT assay indicates reduced NHA cell proliferation by genetic inhibition of the msiCAT-tailing pathway (n=3; unpaired Student’s t-test; **, P < 0.01; ****, P <0.0001, compared to controls at the corresponding time). (D, F) Transwell assay reveals that both genetic (D) and pharmacological (F) inhibition of the msiCAT-tailing pathway hampers the migration of GBM cells but not control cells. (E, G) Quantification of (D, F) showing the number of migrated cells (n=3; unpaired Student’s t-test; ***, P < 0.001; ****, P <0.0001; ns, not significant). (H, J) TUNEL staining shows that both genetic (H) and pharmacological (J) inhibition of the msiCAT-tailing pathway enhances STS-induced apoptosis in GBM cells, using TUNEL-Cy3 as an apoptotic cell indicator and DAPI as a nucleus indicator. (I, K) Quantification of (H, J) showing the percentage of TUNEL-positive cells in the population (n=3; unpaired Student’s t-test; **** P < 0.0001), using DMSO as the vehicle control. (L) MTT assay shows that pharmacological inhibition of the msiCAT-tailing pathway decrease the resistance of GBM cells to temozolomide (TMZ, 150 µM) treatment (n=3; unpaired Student’s t-test; * P < 0.05; **, P < 0.01). (M) Neurosphere formation assay shows that reduced spheroid formation caused by pharmacological inhibition of the msiCAT-tailing pathway can synergize with TMZ in GBM cells. (N) Quantification of (M) (n=3; unpaired Student’s t-test; **, P < 0.01).

Revised version: (A) Cell viability assay shows greater sensitivity to anisomycin treatment in patient-derived GSC cells than control NSC cells at 48 hours (n=3; unpaired Student’s t-test; **, P < 0.001; ****, P <0.0001; compared to controls at the corresponding dose). (B) MTT assay indicates reduced GBM cell proliferation by genetic inhibition of the msiCAT-tailing pathway (n=3; unpaired Student’s t-test; **, P < 0.01; ***, P < 0.001; ****, P <0.0001, compared to controls at the corresponding time). (C) MTT assay indicates reduced NHA cell proliferation by genetic inhibition of the msiCAT-tailing pathway (n=3; unpaired Student’s t-test; **, P < 0.01; ****, P <0.0001, compared to controls at the corresponding time). (D, F) Transwell assay reveals that both genetic (D) and pharmacological (F) inhibition of the msiCAT-tailing pathway hampers the migration of GBM cells but not control cells. (E, G) Quantification of (D, F) showing the number of migrated cells (n=3; unpaired Student’s t-test; ***, P < 0.001; ****, P <0.0001; ns, not significant). (H, J) TUNEL staining reveals that both genetic (H) and pharmacological (J) inhibition of the msiCAT-tailing pathway promote STS-induced apoptosis in GBM cells, utilizing TUNEL-Cy3 as an apoptotic cell marker and DAPI as a nuclear stain. (I, K) Quantification of (H, J) showing the percentage of TUNEL-positive cells in the population (n=3; unpaired Student’s t-test; **** P < 0.0001), using DMSO as the vehicle control. (L) MTT assay shows that pharmacological inhibition of the msiCAT-tailing pathway decrease the resistance of GBM cells to temozolomide (TMZ, 150 µM) treatment (n=3; unpaired Student’s t-test; * P < 0.05; **, P < 0.01). (M) The neurosphere formation assay shows that reduced spheroid formation, caused by pharmacological inhibition of the msiCAT-tailing pathway, can synergize with TMZ in GBM cells. (N) Quantification of (M) (n=3; unpaired Student’s t-test; **, P < 0.01).

Further investigation revealed the RQC pathway’s involvement in GBM cell anti-apoptosis, with initial findings pointing to alterations in mitochondrial functions. Prior studies demonstrated that genetic or pharmacological inhibition of the RQC pathway led to a significant decrease in GBM mitochondrial membrane potential (ΔΨm) (Fig. 2C, D). In GSC cells, anisomycin treatment promoted mitochondrial permeability transition pore (MPTP) opening, an effect not seen in NHA cells (Fig. 3A-D). Consequently, GBM cell lines with genetically or pharmacologically inhibited RQC pathways were more susceptible to STS-induced apoptosis, evidenced by elevated executioner caspase 3/7 activity (Fig. S9A), enhanced PARP-1 cleavage (Fig. S9B, C), increased TUNEL staining (Fig. 5H-K), and flow cytometry analysis (Fig. S9D-G). Notably, general translation inhibition using cycloheximide did not elicit the same apoptotic response (Fig. S9A, D, E). Finally, the RQC pathway was also implicated in temozolomide (TMZ)-induced cell death. Combining anisomycin with TMZ significantly reduced GBM cell survival (Fig. 5L) and effectively inhibited GSC spheroid growth (Fig. 5M, N). In summary, the RQC pathway plays a critical role in multiple aspects of GBM progression, including proliferation, migration, and survival under apoptotic stress.

Discussion

Earlier version: The Ribosome-associated Quality Control (RQC) pathway plays a crucial role in managing aberrant proteins produced during translation. This study focused on understanding the consequences of RQC-mediated modification, specifically the addition of msiCAT tails, to mitochondrial proteins like ATP5α in glioblastoma (GBM) cells. The findings reveal that GBM cells harboring msiCAT-modified ATP5α exhibit a unique metabolic profile. Despite a reduction in ATP synthesis, these cells maintain their mitochondrial membrane potential (ΔΨm), a key factor for cellular function and survival. Furthermore, they demonstrate enhanced cell survival and motility, characteristics associated with increased tumor invasiveness and metastasis. Notably, the presence of msiCAT-modified ATP5α confers resistance to apoptosis triggered by staurosporine (STS), potentially by modulating the mitochondrial permeability transition pore (MPTP), a critical regulator of cell death pathways, as illustrated in Figure 6. These identified traits contribute to an increased aggressiveness of tumors, suggesting that the RQC pathway plays a critical role in cancer cell survival and proliferation. Encouragingly, a recent study also demonstrated the RQC pathway’s involvement in a Drosophila Notch overexpression-induced brain tumor model (56). The findings imply modulating the RQC pathway could serve as a promising complementary strategy to existing chemotherapy regimens. By targeting this specific pathway, therapeutic interventions might effectively disrupt the mechanisms that allow cancer cells to evade apoptosis and sustain their energy production under stress, potentially leading to improved treatment outcomes for patients with GBM and other cancers characterized by similar protein modifications.

Revised version: The Ribosome-associated Quality Control (RQC) pathway plays a crucial role in managing aberrant proteins produced during translation. This study focused on understanding the consequences of RQC-mediated modification, specifically the addition of msiCAT tails, on mitochondrial proteins such as ATP5α in glioblastoma (GBM) cells. The findings reveal that GBM cells harboring msiCAT-modified ATP5α exhibit a unique metabolic profile. Despite a reduction in ATP synthesis, these cells maintain their mitochondrial membrane potential (ΔΨm), a key factor for cellular function and survival. Furthermore, they demonstrate enhanced cell survival and motility, characteristics associated with increased tumor invasiveness and metastasis. Notably, the presence of msiCAT-modified ATP5α confers resistance to apoptosis triggered by staurosporine (STS), potentially by modulating the mitochondrial permeability transition pore (MPTP), a critical regulator of cell death pathways, as illustrated in Figure 6. These identified traits contribute to an increased aggressiveness of tumors, suggesting that the RQC pathway plays a critical role in cancer cell survival and proliferation. Encouragingly, a recent study also demonstrated the RQC pathway’s involvement in a Drosophila model of Notch overexpression-induced brain tumors (56). The findings imply modulating the RQC pathway could serve as a promising complementary strategy to existing chemotherapy regimens. By targeting this specific pathway, therapeutic interventions might effectively disrupt the mechanisms that allow cancer cells to evade apoptosis and sustain their energy production under stress, potentially leading to improved treatment outcomes for patients with GBM and other cancers characterized by similar protein modifications.

Impact of msiCAT-tail modified ATP5α protein on mitochondrial function in GBM cells

Earlier version: In healthy cells, ATP5α protein, encoded by the nuclear genome, is imported into the mitochondrial matrix via the TOM/TIM complex through co-translational import and incorporated into ATP synthase (Left). Conversely, in GBM cells, the CAT-tailed ATP5α protein can either form aggregates near the mitochondrial outer membrane or be imported into the mitochondria. Within the mitochondrial matrix, proteins with shorter CAT-tails readily integrate into ATP synthase, disrupting its functionality. This dysfunction is characterized by a reduced ATP synthesis rate and proton (H+) accumulation, leading to an elevated mitochondrial membrane potential (ΔΨm). These alterations in ATP synthase ultimately trigger malfunction of the MPTP, consequently affecting cell proliferation, migration, and resistance to drug-induced apoptosis (Right).

Revised version: In healthy cells, ATP5α protein, encoded by the nuclear genome, is imported into the mitochondrial matrix via the TOM/TIM complex through co-translational import and incorporated into ATP synthase (Left). Conversely, in GBM cells, the CAT-tailed ATP5α protein can either form aggregates near the mitochondrial outer membrane or be imported into the mitochondria. Within the mitochondrial matrix, proteins with shorter CAT-tails readily integrate into ATP synthase, disrupting its functionality. This dysfunction is characterized by a reduced ATP synthesis rate and proton (H+) accumulation, resulting in an elevated mitochondrial membrane potential (ΔΨm). These alterations in ATP synthase ultimately trigger malfunction of the MPTP, consequently affecting cell proliferation, migration, and resistance to drug-induced apoptosis (Right).

Earlier version: The study of ATP synthase behavior in cancer holds particular importance. During carcinogenesis, ATP synthase frequently relocates to the plasma membrane, where it is termed ectopic ATP synthase (eATP synthase). These eATP synthases exhibit catalytic activity, facilitating ATP production in the extracellular space to foster a favorable tumor microenvironment (57). Research indicates that eATP synthase assembles initially in mitochondria before being transported to the cell surface via microtubules (47). However, the specific type of ATP synthase delivered to the plasma membrane remains unclear. Future investigations into the localization of CAT-tailed eATP synthase may offer valuable insights into this process.

Revised version: The study of ATP synthase behavior in cancer holds particular importance. During carcinogenesis, ATP synthase frequently relocates to the plasma membrane, where it is referred to as ectopic ATP synthase (eATP synthase). These eATP synthases exhibit catalytic activity, facilitating ATP production in the extracellular space to foster a favorable tumor microenvironment (57). Research indicates that eATP synthase assembles initially in mitochondria before being transported to the cell surface via microtubules (47). However, the specific type of ATP synthase delivered to the plasma membrane remains unclear. Future investigations into the localization of CAT-tailed eATP synthase may offer valuable insights into this process.

Earlier version: Multiple mitochondrial proteins in cancer cells can likely undergo CAT-tailing in a similar way. These msiCAT-tailed peptides may have varied impacts on mitochondria and cells due to differences in their base proteins. For instance, CAT-tailed COX4 protein might substantially and directly diminish mitochondrial respiratory efficiency. Examining the individual roles of these proteins is important, as the combined effect of their defects may be crucial in understanding observed mitochondrial changes in cancer. A minor caveat here is that the observed effect of the CAT-tails’ presence primarily stems from artificial CAT-tail sequences with a high threonine content, rather than the endogenous CAT-tail protein. It is possible that other sequence components could lead to different effects (44). A recent study found that ANKZF1 knockdown inhibited GBM progression by causing abnormal protein accumulation in mitochondria (29). This, combined with our data, suggests that balanced ANKZF1 expression and activity are vital for cancer proliferation. Both excess and deficiency may alter cellular adaptability. A minor flaw of that study was the use of a mitochondrial-localized non-stopped GFP protein to induce proteostasis stress and the lack of direct biochemical evidence of CAT-tailed proteins. Our research focuses on endogenous proteins for a detailed analysis of their impact on mitochondria. The rationale is that highly expressed, non-physiological ectopic proteins might cause general proteostasis failure, masking the specific functions of endogenous proteins. Additionally, the studies used different cell lines. GSC, a patient-derived GBM cell line with greater stemness, might have distinct mitochondrial status and RQC pathway activity compared to U87 or U251 cell lines. Thus, the conclusions of the two studies are not contradictory but rather complementary, both demonstrating the significance of RQC in tumorigenesis. Our study delves into the mechanistic role of the RQC pathway in GBM, identifying new potential targets for future treatments.

Revised version: Multiple mitochondrial proteins in cancer cells can likely undergo CAT-tailing in a similar way. These msiCAT-tailed peptides may have varied impacts on mitochondria and cells due to differences in their base proteins. For instance, CAT-tailed COX4 protein might substantially and directly diminish mitochondrial respiratory efficiency. Examining the individual roles of these proteins is important, as the combined effect of their defects may be crucial in understanding observed mitochondrial changes in cancer. A minor caveat here is that the observed effect of the CAT-tails’ presence primarily stems from artificial CAT-tail sequences with a high threonine content, rather than the endogenous CAT-tail protein. It is possible that other sequence components could lead to different effects (44). A recent study found that ANKZF1 knockdown inhibited GBM progression by causing abnormal protein accumulation in mitochondria (29). This, combined with our data, suggests that balanced ANKZF1 expression and activity are vital for cancer proliferation. Both excess and deficiency may alter cellular adaptability. A minor flaw of that study was the use of a mitochondrial-localized non-stopped GFP protein to induce proteostasis stress and the lack of direct biochemical evidence of CAT-tailed proteins. Our research focuses on endogenous proteins for a detailed analysis of their impact on mitochondria. The rationale is that highly expressed, non-physiological ectopic proteins might cause general proteostasis failure, masking the specific functions of endogenous proteins. Additionally, the studies used different cell lines. GSC, a patient-derived GBM cell line with greater stemness, might have a distinct mitochondrial status and RQC pathway activity compared to U87 or U251 cell lines. Thus, the conclusions of the two studies are not contradictory but rather complementary, both demonstrating the significance of RQC in tumorigenesis. Our study delves into the mechanistic role of the RQC pathway in GBM, identifying new potential targets for future treatments.

An in-depth investigation into the quantification of nuclear genome-encoded mitochondrial proteins modified via the msiCAT-tailing mechanism using sophisticated mass spectrometry is a compelling area for future research. Recent work by Lv et al., published in Cell Reports, revealed that the cytoplasmic E3 ligase Pirh2 and the mitochondrial protease ClpXP work in conjunction with the established NEMF-ANKZF1 system to break down mitochondrial protein aggregates resulting from ribosome stalling (58). The increased presence of ClpXP in various cancers could potentially be linked to an increase in msiCAT-tailing products in mitochondria, though further studies are needed to clarify ClpXP’s role in mitochondrial RQC (59). Moreover, ClpXP influences the levels of multiple mitochondrial proteins. Our own experiments showed that ATP5α proteins lacking msiCAT-tails were the most challenging to express ectopically. Proteins with shorter tails (AT3) expressed more readily, while those with longer tails (AT20) exhibited the highest expression levels but also tended to form SDS-insoluble aggregates. This regulatory effect could be mediated by ClpXP-dependent degradation or potentially through transcriptional control. PGC-1α, the peroxisome proliferator-activated receptor gamma co-activator, is a key regulator of mitochondrial biogenesis in mammals (60). By binding to and activating nuclear transcription factors, PGC-1α triggers the transcription of nuclear genome-encoded mitochondrial proteins and the mitochondrial transcription factor Tfam. Tfam, in turn, activates mitochondrial genome transcription and replication (61). Distinguishing between these regulatory possibilities will necessitate future research, including a meticulous examination of mRNA levels for msiCAT-tailed targets and analysis of PGC-1α and Tfam binding to transcriptional elements.

MPTP is a complex, supramolecular channel traversing the inner mitochondrial membrane, characterized by its non-selective ion permeability, calcium dependence, and multifaceted functionality. Despite extensive investigations into its functional attributes and regulatory mechanisms, the precise molecular architecture of the MPTP remains elusive (62). Several theoretical models have been posited to elucidate the MPTP’s structural composition. Firstly, the VDAC/ANT/Cyp-D model (63) proposed an assembly of voltage-dependent anion channels (VDAC), adenine nucleotide translocators (ANT), and cyclophilin D (Cyp-D) as the structural basis; however, subsequent genetic analyses have introduced substantial controversy regarding the integral role of these proteins within the MPTP complex (64–67). Secondly, the ATP synthase model posits that MPTP formation involves dimers or reconstituted c-rings of ATP synthase (48, 49). While this hypothesis presents an intriguing perspective, empirical confirmation of ATP synthase’s role as a definitive structural element of the pore remains inconclusive, with a body of conflicting research surrounding this proposition. Thirdly, the contemporary prevailing hypothesis suggests the MPTP is constituted by a large complex, termed the ATP synthasome, comprising ANT and ATP synthase, with Cyp-D serving a regulatory function over the complex’s dynamic behavior (68).

The MPTP activity is modulated by mitochondrial membrane potential (ΔΨm), which reciprocally influences mitochondrial ion homeostasis and energy metabolism (69, 70). Our study elucidates a dual function of msiCAT-tailed ATP5α protein in cancer cells: stabilization of a high membrane potential, thereby mitigating MPTP induction, and direct inhibition of MPTP functionality through participation in its assembly. While MPTP’s critical role in cell death is established, the premise that MPTP inhibition enables cancer cell evasion of drug-induced programmed cell death has lacked substantial evidence. This study furnishes empirical support for this hypothesis, demonstrating that GBM cells, notably glioblastoma stem cells (GSC), exhibit markedly reduced MPTP activity relative to control cells. This reduced activity is directly correlated with the CAT-tailing modification of the ATP synthase subunit. These observations are concordant with prior research indicating that genetic mutations or post-translational modifications in specific ATP synthase subunits can modulate MPTP activity. The findings highlight a novel mechanism through which cancer cells may develop resistance to therapeutic interventions by manipulating mitochondrial function (71, 72).

Added in the revised version: Methods

Removed in the revised version: Materials and methods

Cell lines and cell culture conditions

The human astroglia cell line SVG p12 (ATCC, cat. CRL-8621) and the human glioma cell line SF268 were from Dr. Rongze Olivia Lu. Both cell lines were cultured in DMEM (ATCC, cat. #302002) with 10% FBS (Biowest, cat. S1620-100) and penicillin/streptomycin (Gibco™, cat. 15140122). SF268 clones should be maintained in complete DMEM supplemented with 400 µg/mL G418 (Gibco, cat. 10131027). The 0.25% trypsin solution (ATCC, cat. #SM2003C) was used to passage cells. The normal human astrocytes NHA E6/E7/hTERT cell line was from Dr. Russell O. Pieper, UCSF Brain Tumor Research Center. Cells are cultured in ABMTM Basal Medium (Lonza, cat. CC-3187) and AGMTM SingleQuotsTM Supplements (Lonza, cat. CC-4123). Corning™ Accutase™ Cell Detachment Solution (Corning, cat. 25058CI) was used to passage cells. GSC827, a patient-derived human glioma stem cell line, was from Dr. Chun-Zhang Yang at NIH. The NSC, NSC26, patient-derived GSC33, GSC22, GSC99, GSC105, and GSC107 cell lines used in this study were kindly provided by Dr. John S Kuo at the University of Texas, Austin. GSC cells were cultured in Neural basal-A Medium (Gibco, cat. #10888022) with 2% B27 (Gibco, cat. #17504044), 1% N2 (Gibco, cat. #17502048), 20 ng/ml of EGF and FGF (Shenandoah Biotechnology Inc. cat. PB-500-017), Antibiotic-Antimycotic (Gibco, cat. #15240062), and L-Glutamine (Gibco, cat. #250300810). Cells could be cultured in both spherical and attached (on Geltrex, Thermo Fisher, cat. A1413202) forms. Corning™ Accutase™ Cell Detachment Solution (Corning, cat. 25058CI) was used to passage cells.

Cells were transfected with X-tremeGENE™ HP DNA Transfection Reagent (Sigma, cat. 6366244001) following the standard protocol. For single clone selection, SF268 cells were treated with 800 µg/ml G418 for 5 days. The cells were then seeded into a 96-well plate at a density of 1/100 µL. Positive clones were verified by immunofluorescence staining and immunoblotting. Cells were maintained in complete DMEM containing 400 µg/mL G418. GBM cell lines were subjected to a 4-hour pre-treatment at 37°C using either anisomycin (20 nM or 200 nM, Fisher Scientific, cat. AAJ62964MF) or cycloheximide (100 µg/mL, Fisher Scientific, cat. AC357420010) in medium, as detailed in the conducted experiments.

Primers, plasmids and viruses

Plasmids pcDNA3.1+/C-(K)-DYK-ATP5F1A (pATP511 control), pcDNA3.1+/C-(K)-DYK-ATP5F1A-AT3 (pATP511-AT3), pcDNA3.1+/C-(K)-DYK-ATP5F1A-AT20 (pATP511-AT20), pcDNA3.1+/C-(K)-DYK-ATP5F1A-GS3 (pATP511-GS3), and pcDNA3.1+/C-(K)-ATP5F1A-DYK-GS20 (pATP511-GS20) were generated by GenScript Inc. Plasmids pCMV-5×FLAG-β-globin-control (5FBG-Ctrl) and pCMV-5×FLAG-β-globin-non-stop (5FBG-nonstop) were generated by Dr. Hoshino (Nagoya City University) and Dr. Inada (Tohoku University) (43). pCMV6-DDK-NEMF (oeNEMF) was from ORIGENE Inc. (cat. RC216806).

Viruses (and plasmid used to generate viruses) are pLV[CRISPR]-hCas9:T2A:Neo-U6>Scramble[gRNA#1] (sgControl/sgCtrl), pLV[CRISPR]-hCas9:T2A:Neo-U6>hNEMF[gRNA#1579] (sgNEMF), pLV[Exp]-Bsd-EF1A>ORF_Stuffer (pLV-control), pLV[Exp]-EGFP:T2A:Puro-EF1A>mCherry (pLV-control-2/oeCtrl), pLV[Exp]-Bsd-EF1A>hANKZF1[NM_001042410.2]/HA (oeANKZF1), and pLV[Exp]-mCherry/Neo-EF1A>hANKZF1[NM_001042410.2] (oeANKZF1) were made by VectorBuilder Inc.

Primers (5’ to 3’) used for RT-PCR are:

Neurosphere formation assay of GSCs

Earlier version: The GSC spheroids were dissociated using Accutase for 2 min. Cells were resuspended in a single-cell suspension and grown under non-adherent conditions. Cells were seeded in 12-well plates at a density of 0.25x106 cells/well and cultured in 3 mL culture medium for 24 hours. 20 nM of anisomycin and 150 µM of temozolomide (TMZ) were added to the culture medium and treated for 96 hours. Spheroids were imaged under a 10x objective, captured using QCapture, and analyzed with ImageJ. Spheroids larger than 50 µm were counted.

Revised version: The GSC spheroids were dissociated using Accutase for 2 min. Cells were resuspended in a single-cell suspension and grown under non-adherent conditions. Cells were seeded in 12-well plates at a density of 0.25 x 106 cells/well and cultured in 3 mL culture medium for 24 hours. 20 nM of anisomycin and 150 µM of temozolomide (TMZ) were added to the culture medium and the cells were treated for 96 hours. Spheroids were imaged under a 10x objective, captured using QCapture, and analyzed with ImageJ. Spheroids larger than 50 µm were counted.

Differential gene expression analysis using the public database

Earlier version: The raw RNA-seq data used to perform the analysis were obtained from the University of California, Santa Cruz Xenabrowser (cohort: TCGA TARGET GTEx, dataset ID: TcgaTargetGtex_rsem_gene_tpm, https://xena.ucsc.edu/), and then subsets included only TCGA glioma (GBM), GTEx Brain Frontal Cortex, and GTEx Cortex samples. Differential expression analysis was conducted using the “Limma” package (R version: 4.3.1). The Benjamini-Hochberg method was used for multiple testing correction to control the false discovery rate (FDR). Cut-off of adjusted p-value (adj.P.Val) was set at 0.001; cut-off of the absolute fold change was set at 2 (logFC > 1).

Revised version: The raw RNA-seq data used for the analysis were obtained from the University of California, Santa Cruz Xenabrowser (cohort: TCGA TARGET GTEx, dataset ID: TcgaTargetGtex_rsem_gene_tpm, https://xena.ucsc.edu/). Subsets were then created to include only TCGA glioma (GBM), GTEx Brain Frontal Cortex, and GTEx Cortex samples. Differential expression analysis was conducted using the “Limma” package (R version: 4.3.1). The Benjamini-Hochberg method was used for multiple testing correction to control the false discovery rate (FDR). Cut-off of adjusted p-value (adj.P.Val) was set at 0.001; cut-off of the absolute fold change was set at 2 (logFC > 1). The code is available without restrictions at https://github.com/yuanna23/GBM_elife.

Immunostaining

Earlier version: Cells were cultured on sterile coverslips until 80% confluency. For immunostaining, cells were washed with phosphate-buffered saline (PBS) solution thrice. Then, 4% formaldehyde (Thermo Fisher, cat. BP531-500) was applied to cells for fixation for 30 min at room temperature. After fixation, cells were washed with PBS solution containing 0.25% Triton X-100 (PBSTx) (Thermo Fisher, cat. T9284) thrice, and blocked with 5% normal goat serum (Jackson Immuno, cat. 005-000-121) for 1 hour at room temperature. Cells were then incubated with primary antibodies overnight in a humidified chamber at 4°C. The next day, cells were washed by PBSTx thrice and incubated with secondary antibodies for 2 hours at room temperature. After washing, cells were stained with 300 nM DAPI (Thermo Fisher, cat. 57-481-0) for 5 min at room temperature and mounted in Fluoromount-G Anti-Fade solution (Southernbiotech, cat. 0100-35). Images were taken using a Zeiss LSM 800 confocal microscope, 40x oil objective lens and Airyscan processing. The primary antibodies used in the study were rabbit anti-ATP5a (Cell Signaling, cat. #18023), mouse anti-TOMM20 (1:500, Santa Cruz, cat. 18023S), rabbit anti-MTCO2 (1:500, Abclonal, cat. sc-17764), mouse anti-NDUS3 (1:1000, Abcam, cat. ab14711). The secondary antibodies were Alexa fluor 633-, 594-, 488-conjugated secondary antibodies (1:300, Invitrogen, cat. A21071, A11036, A32732).

Revised version: Cells were cultured on sterile coverslips until 80% confluency. For immunostaining, cells were washed with phosphate-buffered saline (PBS) solution thrice. Then, 4% formaldehyde (Thermo Fisher, cat. BP531-500) was applied to cells for fixation for 30 min at room temperature. After fixation, cells were washed with PBS solution containing 0.25% Triton X-100 (PBSTx) (Thermo Fisher, cat. T9284) thrice, and blocked with 5% normal goat serum (Jackson Immuno, cat. 005-000-121) for 1 hour at room temperature. Cells were then incubated with primary antibodies overnight in a humidified chamber at 4°C. The next day, cells were washed by PBSTx thrice and incubated with secondary antibodies for 2 hours at room temperature. After washing, cells were stained with 300 nM DAPI (Thermo Fisher, cat. 57-481-0) for 5 min at room temperature and mounted in Fluoromount-G Anti-Fade solution (Southern Biotech, cat. 0100-35). Images were captured using a Zeiss LSM 800 confocal microscope with a 40x oil objective lens and AiryScan processing. The primary antibodies used in the study were rabbit anti-ATP5a (Cell Signaling, cat. #18023), mouse anti-TOMM20 (1:500, Santa Cruz, cat. sc-17764), rabbit anti-MTCO2 (1:500, Proteintech, cat. 55070-1-AP), and mouse anti-NDUS3 (1:1000, Abcam, cat. ab14711). The secondary antibodies were Alexa fluor 633-, 594-, 488-conjugated secondary antibodies (1:300, Invitrogen, cat. A21071, A11036, A32732).

SDS-PAGE and immunoblotting

Cells or isolated mitochondria were solubilized in cell lysis buffer containing 50 mM Tris-HCl pH 7.4, 150 mM NaCl, 10% glycerol, 1% Triton X-100, 5 mM EDTA, and 1x protease inhibitor (Bimake, cat. B14002). Protein concentration was measured by using the Bradford assay (BioVision, cat. K813-5000-1). Samples were separated in a 4-12% Tris-Glycine gel (Invitrogen, cat. WXP41220BOX) and proteins were transferred to a PVDF membrane (Millipore, cat. ISEQ00010). The membranes were then blocked with 5% non-fat dry milk (Kroger) for 50 min at room temperature and probed with primary antibodies overnight at 4°C. Membranes were washed with Tris-buffered saline with 0.1% Tween 20 (TBST) solution thrice and then incubated with secondary antibodies for 1 hour at room temperature. Blots were detected with ECL solution (PerkinElmer, cat. NEL122001EA) and imaged by Chemidoc system (BioRad). The intensity of blots was further analyzed by ImageJ software. The primary antibodies used were mouse anti-Actin (1:1000, Santa Cruz, cat. sc-47778), rabbit anti-NEMF (1:1000, Proteintech, cat. 11840-1-AP), mouse anti-ANKZF1 (1:1000, Santa Cruz, cat. sc-398713), mouse anti-ATP5a (Abcam, cat. ab14748), mouse anti-NDUS3 (1:1000, Abcam, cat. ab14711), rabbit anti-COX4 (Abcam, cat. ab209727), mouse anti-Flag (1:1000, Sigma, cat. F1804), rabbit anti-ANT1/2 (1:1000, Proteintech, cat. 17796-1-AP), rabbit anti-CypD (1:1000, Proteinetch, cat. 15997-1-AP), rabbit anti-PARP1 (1:1000, Abclonal, cat. A0942), rabbit anti-GAPDH (1:1000, Abclonal, cat. A19056). The secondary antibodies used were goat anti-rabbit IgG (1:5000, Invitrogen, cat. G21234), goat anti-mouse IgG (1:5000, Invitrogen, cat. PI31430).

Mitochondrial isolation, blue Native PAGE, and western blotting

Earlier version: Cells were homogenized using Dounce homogenizer in ice-cold homogenization buffer containing 210 mM mannitol (Fisher Sci, cat. AA3334236), 70 mM sucrose (Fisher Sci, cat. AA36508A1), 5 mM HEPES (Fisher Sci, cat. 15630106), pH 7.12, 1 mM EGTA (Fisher Sci, cat. 28-071-G), and 1x protease inhibitor. The homogenate was centrifuged at 1500 g for 5 min. The resultant supernatant was centrifuged at 13000 g for 17 min. The supernatant was collected as the cytosol portion, and the pellet as the mitochondria portion, was washed with homogenization buffer and centrifuged at 13000 g for 10 min. For Blue Native PAGE, the mitochondria samples were solubilized in 5 % digitonin (Thermo Fisher, cat. BN2006) on ice for 30 min and then centrifuged at 20000 g for 30 min. The supernatant contains solubilized mitochondrial proteins and was mixed with 5% G-250 (GoldBio, cat. C-460-5) and 1x NativePAGE sample buffer (Invitrogen, cat. BN2008) (final G-250 concentration is 25% of the digitonin concentration). Mitochondrial protein concentration was measured by using the Bradford assay. Samples were separated in 3-12% Bis-Tris Native gel (Invitrogen, cat. BN1001BOX) and then transferred to a PVDF membrane. Membranes were fixed with 8% acetic acid (Thermo Fisher, cat. 9526-33), and then blocked and probed with antibodies as described above for Western blotting.

Revised version: Cells were homogenized using Dounce homogenizer in ice-cold homogenization buffer containing 210 mM mannitol (Fisher Sci, cat. AA3334236), 70 mM sucrose (Fisher Sci, cat. AA36508A1), 5 mM HEPES (Fisher Sci, cat. 15630106), pH 7.12, 1 mM EGTA (Fisher Sci, cat. 28-071-G), and 1x protease inhibitor. The homogenate was centrifuged at 1500 g for 5 min. The resultant supernatant was centrifuged at 13000 g for 17 min. The supernatant was collected as the cytosol portion, and the pellet (the mitochondria portion) was washed with homogenization buffer and centrifuged at 13000 g for 10 min. For Blue Native PAGE, the mitochondrial samples were solubilized in 5% digitonin (Thermo Fisher, cat. BN2006) on ice for 30 min and then centrifuged at 20,000 g for 30 min. The supernatant contains solubilized mitochondrial proteins and was mixed with 5% G-250 (GoldBio, cat. C-460-5) and 1x NativePAGE sample buffer (Invitrogen, cat. BN2008) (final G-250 concentration is 25% of the digitonin concentration). Mitochondrial protein concentration was measured by using the Bradford assay. Samples were separated in 3-12% Bis-Tris Native gel (Invitrogen, cat. BN1001BOX) and then transferred to a PVDF membrane. Membranes were fixed with 8% acetic acid (Thermo Fisher, cat. 9526-33), and then blocked and probed with antibodies as described above for Western blotting.

Mitochondrial membrane potential assays

Earlier version: Mitochondrial membrane potential of GSC cells was measured using Image-iTTM TMRM (Invitrogen, cat. I34361). Cells were cultured in the 96-well black plate at a density of 1 x105 cells per well overnight in an incubator with 5% CO2 at 37°C. Cells were incubated with TMRM (100 nM) for 30 min at 37°C. Then, cells were washed with PBS solution three times. Fluorescence changes at excitation/emission of 548/574 nm were monitored with a Cytation 5 plate reader (BioTek). Mitochondrial membrane potential was also measured using JC-10 (AdipoGen, cat. 50-114-6552). Cells were cultured in the 96-well black plate at a density of 5 x 104 cells per well overnight in an incubator with 5% CO2 at 37°C. Cells were incubated with JC-10 (10 µg/ml) for 45 min at 37°C. Then, cells were washed with PBS solution twice. Fluorescence changes at excitation/emission of 535/595 nm for JC-10 aggregates and at 485/535 nm for JC-10 monomers were monitored with a Synergy 2 Reader (BioTek). Mitochondrial membrane potential was quantified as the fluorescence of JC-10 aggregates/monomers (595/535 nm).

Revised version: Mitochondrial membrane potential of GSC cells was measured using Image-iTTM TMRM (Invitrogen, cat. I34361). Cells were cultured in 96-well black plates at a density of 1 x 105 cells per well overnight in an incubator with 5% CO2 at 37°C. Cells were incubated with TMRM (100 nM) for 30 min at 37°C. Then, cells were washed with PBS solution three times. Fluorescence changes at excitation/emission of 548/574 nm were monitored with a Cytation 5 plate reader (BioTek). Mitochondrial membrane potential was also measured using JC-10 (AdipoGen, cat. 50-114-6552). Cells were cultured in 96-well black plates at a density of 5 x 104 cells per well overnight in an incubator with 5% CO2 at 37°C. Cells were incubated with JC-10 (10 µg/ml) for 45 min at 37°C. Then, cells were washed with PBS solution twice. Fluorescence changes at excitation/emission of 535/595 nm for JC-10 aggregates and at 485/535 nm for JC-10 monomers were monitored with a Synergy 2 Reader (BioTek). Mitochondrial membrane potential was quantified as the fluorescence of JC-10 aggregates/monomers (595/535 nm).

Seahorse cell mitochondrial stress assays

Earlier version: The oxygen consumption rate (OCR) of cells was measured using the Seahorse Cell Mito Stress Test kit following the user guide (Agilent, cat. 103010-100). Briefly, cells were cultured in testing chambers at a density of 8000 cells per well overnight in an incubator with 5% CO2 at 37°C. Cells were then washed with assay medium containing Seahorse XF DMEM medium (Agilent, cat. 103575-100) with 1 mM pyruvate, 2 mM glutamine, and 10 mM glucose twice, and incubated in assay medium for 1 hour in an incubator without CO2 at 37°C. Cells were treated with compounds in the order of oligomycin (1.5 µM), carbonyl cyanide-4 (trifluoromethoxy), phenylhydrazone (FCCP, 1.0 µM), and Rotenone/Antimycin (0.5 µM). The OCR of cells was monitored by using Seahorse XF HS Mini (Agilent).

Revised version: The oxygen consumption rate (OCR) of cells was measured using the Seahorse Cell Mito Stress Test kit following the user guide (Agilent, cat. 103010-100). Briefly, cells were cultured overnight in testing chambers at a density of 8,000 cells per well in an incubator with 5% CO2 at 37°C. Cells were then washed twice with assay medium containing Seahorse XF DMEM medium (Agilent, cat. 103575-100) supplemented with 1 mM pyruvate, 2 mM glutamine, and 10 mM glucose. They were subsequently incubated in the assay medium for 1 hour in an incubator without CO2 at 37°C. Cells were treated with compounds in the order of oligomycin (1.5 µM), carbonyl cyanide-4 (trifluoromethoxy), phenylhydrazone (FCCP, 1.0 µM), and Rotenone/Antimycin (0.5 µM). The OCR of cells was monitored by using Seahorse XF HS Mini (Agilent).

Mitochondrial MPTP assay

The status of mitochondrial permeability transition pore was measured using Invitrogen™ Image-IT™ LIVE Mitochondrial Transition Pore Assay Kit (Invitrogen, cat. I35103). Cells were cultured in 35 mm glass-bottom dishes overnight in an incubator with 5% CO2 at 37°C. Cells were washed twice with the modified Hank’s Balanced Salt Solution (HBSS, Thermo Fisher, cat. 14025092) containing 10 mM HEPES, 2 mM L-glutamine and 0.1 mM succinate (Thermo Fisher, cat. 041983.A7) and incubated with the labeling solution (1 µM Calcein, 0.2 µM MitoTracker Red, 1 mM Cobalt Chloride) for 15 min at 37°C. Cells were then washed with HBSS twice and imaged at excitation/emission of 494/517 nm for Calcein and at 579/599 nm for MitoTracker Red by using the Zeiss confocal microscope.

Mitochondrial Ca2+ retention capacity assay

The mitochondrial calcium retention capacity (CRC) was measured on a Cytation 5 reader at excitation/emission of 506/592 nm using the membrane-impermeable fluorescent probe Calcium green-5N (Invitrogen, cat. C3737). Isolated mitochondria samples (0.75 mg protein/mL) were incubated in 1 mL swelling medium supplemented with 10 mM succinate, 1 μM Calcium green-5N, inorganic phosphate, and cyclosporine A (Thermo Fisher, cat. AC457970010). One Ca2+ addition was 1.25 nmol (1 mL volume). Only the MPTP opening in the presence of cyclosporine A was induced by high amounts of added calcium (30 nmol Ca2+ in the last two additions). The CRC value was calculated as total Ca2+ accumulated in the mitochondria per unit (1 mg protein).

MTT assay

Cell proliferation was measured by using the MTT assay kit (Roche, cat. 11465007001). Cells were cultured in 96-well plates at a density of 2000 cells per well overnight in an incubator with 5% CO2 at 37°C. Cells were treated with MTT labeling reagent for 4 hours at 37°C. The solubilization buffer was added to the cells, and then the cells were incubated overnight at 37°C. Absorbance changes of the samples at 550 nm were monitored by using a Synergy 2 Reader (BioTek).

Wound healing assay

Cells were seeded into 6-well plates and cultured for 24-48 hours to reach a confluent cell monolayer. Cells were treated with serum-free medium overnight before mechanical scratching (54). Images of the wounds were taken at 0, 24, and 48 hours. Wound areas were measured by using the wound healing plugin of ImageJ. Wound Coverage % = 100% x [At=0h-At=Δh]/At=0h(At=0h is the area of the wound measured immediately after scratching t = 0h, At=Δh is the area of the wound measured h hours after the scratch is performed).

Cell migration assay

Cell migration was measured by using Transwell assays (Corning, cat. CLS3422). Cells were cultured in Transwell inserts at a density of 1 x 105 cells per well for 3 hours in an incubator at 37°C with 5% CO2. The top inserts were supplemented with DMEM medium only, and the bottom wells were supplemented with DMEM medium with 20% Fetal Bovine Serum. After incubation, the cells on the apical side of the Transwell insert membrane were removed using a cotton applicator. The cells on the bottom side of the insert were rinsed with PBS twice and fixed in 70% ethanol (Thermo Fisher, cat. R40135) for 15 min at room temperature. After fixation, inserts were placed into an empty well to allow the membrane to dry. Then, the insert was incubated with 0.2% crystal violet (Sigma, cat. V5265) for 5 min at room temperature. The insert was rinsed with water twice, and images were captured by using a microscope with a 20x objective. Cell numbers were quantified using ImageJ.

TUNEL staining

The apoptosis was measured by a TUNEL assay kit (ApexBio, cat. K1134). Cells were cultured on sterile cover slips until 80% confluency and washed with PBS thrice. Then, 4% formaldehyde was applied to cells and fixed for at 4°C 25 min. After fixation, cells were washed with PBS twice and incubated with 20 µM proteinase K (Invitrogen, cat. 25530049) for 5 min at room temperature. Then, cells were rinsed with PBS thrice and incubated in 1x equilibration buffer for 10 min at room temperature. Cells were stained with FITC or Cy3 labeling mix for 1 hour at 37°C in a humidified chamber. Cells were washed by PBS thrice and stained with DAPI for 5 min at room temperature. Cells were mounted in the Fluoromount-G Anti-Fade solution and imaged at 520 nm for FITC or at 570 nm for Cy3 by using the Zeiss confocal microscope.

Caspase-3/7 activity assay

Caspase-3/7 activity was measured by using CellEvent™ Caspase 3/7 Detection Reagents (Invitrogen, cat. C10432) following the manufacturer’s protocol. Specifically, cells were seeded in a 96-well black plate with a clear bottom at a density of 5 x 104 cells per well and incubated overnight in the incubator with 5% CO2 at 37°C. Cells were then incubated with 1x staining solution for 30 min at 37°C. Fluorescence changes at excitation/emission of 485/525 nm were monitored with a Synergy 2 Reader (BioTek).

Annexin V-FITC/Propidium Iodide (PI) apoptosis detection

Annexin V-FITC/PI apoptosis assay was performed by using the FITC Annexin V Apoptosis Detection Kit with PI (BioLegend, cat. 640914). Briefly, 1 x 105 cells were collected in 100 μL of staining buffer. Then, cells were incubated with 5 μL of Annexin V-FITC and 2.5 μL of PI for 15 min at room temperature in the dark. Following incubation, 400 μL of binding buffer was added to the stained cells. Flow cytometry analysis of the fluorescence was performed using a Soni SH800 Cell Sorter.

Mitochondria ATP measurement via fluorescence imaging of ATP-red

BioTracker™ ATP red dye (Millipore, cat. SCT045) is a fluorogenic indicator for ATP in mitochondria (73). Cells cultured in monolayer conditions were incubated in medium with 5 μM ATP red for 15 min in an incubator at 37°C with 5% CO2. Mitochondria were also labeled by incubating cells with 100 nM MitoTracker-Green (Invitrogen, cat. M7514) for 15 min to normalize their mass. Before measurement, cells were washed twice with culture medium, and then fresh medium was added. Cells were imaged in a 37°C chamber with 5% CO2 at excitation/emission of 510/570 nm for ATP-red and at excitation/emission of 490/516 nm for MitoTracker-Green by using the Zeiss confocal microscope. The ATP-red signals could also be measured by a Synergy 2 Reader (BioTek).

Co-immunoprecipitation

Cells were lysed in the buffer containing 50 mM Tris-HCl pH 7.4, 150 mM NaCl, 10% glycerol, 1% TritonX-100, 5 mM EDTA, and 1x protease inhibitor. Soluble samples were incubated with 1.5 µL ATP511 antibody at 4°C with mixing overnight. 25 µL of protein A/G magnetic beads (Pierce, cat. 88802) were added to the co-IP samples and incubated at 4°C with mixing overnight. Samples were washed with washing buffer thrice and then applied to SDS-PAGE analysis.

Mice and immunostaining

Animal studies were approved by the University of California, San Francisco Institutional Animal Care and Use Committee (IACUC, AN195636-01) and were performed following the guidelines of the National Institutes of Health (NIH).

For orthotopic brain tumor models, 8-to-10-week-old C57BL/6 mice (male and female in equal numbers) were used for i.c. studies. Cell lines (GL261, SB28) were suspended in DMEM for inoculation. Mice were anesthetized with isoflurane, and 30,000 tumor cells were injected orthotopically in 3 μL. Using a stereotactic frame, a burr hole was formed on the skull via a 0.7 mm drill bit 1.5mm laterally to the right and 1.5mm rostrally from the bregma, and a noncoring needle (26s gauge; Hamilton) was used to inject the cells at a depth of 3mm into the brain from the burr hole. The skin incision was sutured. Mice were then monitored daily. Mouse SB28 tumor tissue and wild-type mouse brain tissue were collected at the survival endpoint.

Frozen tissue sections were thawed at room temperature for 20 min and rinsed with PBS three times. Tissues were then fixed in 4% formaldehyde for 15 min at room temperature. After washing in PBS, tissues were permeabilized with 0.01% Triton X-100 + 0.1% Tween-20 for 15 min and then blocked by using 5% normal goat serum and M.O.M. blocking reagent (Vector Laboratories, cat. BMK-2202) for 1 hour at room temperature. Tissues were then incubated with primary antibodies overnight in a humidified chamber at 4°C. After washing in PBST, tissues were incubated with secondary antibodies for 1 hour. After washing again in PBST, tissues were stained with 300 nM DAPI for 5 min and mounted in Fluoromount-G Anti-Fade solution. Images were taken using a Zeiss LSM 800 confocal microscope. The primary antibodies used in the study were mouse anti-ATP5a (1:500, Abcam, cat. Ab14748), rat anti-TOMM20 (1:500, Abcam, cat. Ab289670), rabbit anti-NEMF (1:500, Proteintech, cat. 11840-1-AP), mouse anti-ANKZF1 (1:500, Santa Cruz, cat. sc-398713), chicken anti-GFP (1:500, Abcam, cat. Ab13970). The secondary antibodies were Alexa Fluor 633-, 594-, 488-conjugated secondary antibodies (1:300, Invitrogen, cat. A21071, A11036, A32732).

Statistics

Earlier version: Statistical analyses were performed by using GraphPad 7.0 software. Chi-squared test and unpaired Student’s t-test were used for comparison. P < 0.05 was considered significant, except in gene expression analysis (Fig. 1A). *, P < 0.05; **, P < 0.01; ***, P < 0.001; ****, P < 0.0001; ns, not significant. All data were expressed as means ± s.e.m.

Revised version: Statistical analyses were performed using GraphPad Prism 9.4. Chi-squared test and unpaired Student’s t-test were used for comparison. P < 0.05 was considered significant, except in gene expression analysis (Fig. 1A). *, P < 0.05; **, P < 0.01; ***, P < 0.001; ****, P < 0.0001; ns, not significant. All data were expressed as means ± s.e.m. This study’s replicates, samples, groups, and experiments were biologically independent, except in Table 1. The “n” numbers for each assay are indicated in the figure legends.

Added in the revised version: Materials availability

Added in the revised version: Plasmids and other reagents generated in this study will be made available to researchers by contacting zhihaowu@smu.edu.

Added in the revised version: Code availability

Added in the revised version: The code used for differential gene expression analysis is available without restrictions at https://github.com/yuanna23/GBM_elife.